```python
# this is needed for the plotting questions.
# if this fails, you may need to run jupyter
# with administrator privileges
import sys
!{sys.executable} -m pip install plotchecker
```

```python
from autograder_term2week3 import *
_globals = globals()
import numpy as np
import matplotlib.pyplot as plt
%matplotlib inline
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki</center>

# <center> Worksheet 7: Yet More About Functions, and numpy</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

## Question 1

(a) Write and test a function called <code>pop_min</code> that takes as its argument a non-empty list <code>num_list</code>, assumed to consist of ints and/or floats, and outputs its minimum value. However, your function should have the side-effect of removing this minimum value from <code>num_list</code>. 

Test your function. If you type

```python
this_list = [3, 5, 2, 7, 1, 2, 3, 1, 9]
pop_min(this_list)
```

then you should get the output 1, and as a side-effect, the value of <code>this_list</code> should become

<code>[3, 5, 2, 7, 2, 3, 1, 9]</code>.

<!-- grade_id: cell-1ae855be0b82fa0e, points: 0, grade: false, solution: true -->
```python
def pop_min(num_list):
    ### BEGIN SOLUTION
    # initialize min at infinity
    this_min = num_list[0]
    min_index = 0
    # for loop
    for i in range(1,len(num_list)):
        # adjust minimum if necessary
        if num_list[i] < this_min:
            min_index = i
            this_min = num_list[i]
    return num_list.pop(min_index)
    ### END SOLUTION
```

<!-- grade_id: cell-31d6d5b9b62d0a0c, points: 0, grade: true, solution: false -->
```python
# Function testing; not for marks
this_list = [3, 5, 2, 7, 1, 2, 3, 1, 9]
this_min = pop_min(this_list)
try:
    assert(this_min == 1)
    assert(this_list == [3, 5, 2, 7, 2, 3, 1, 9])
    print('Test passed!')
except:
    print('Test failed!')
```

<!-- grade_id: cell-4b1a182e62d8ffcf, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

(b) Write and test a function called <code>simple_sort</code> that takes as its argument a list <code>num_list</code>, assumed to be numerical, and, by using <code>pop_min</code>, returns all the elements of <code>num_list</code>, in ascending order. 

<!-- grade_id: cell-544ce6138c69ebf5, points: 0, grade: false, solution: true -->
```python
def simple_sort(num_list):
    ### BEGIN SOLUTION
    # length of list
    n = len(num_list)
    # empty new list
    new_list = []
    # for loop
    for r in range(n):
        # pop the minimum and append it to new list
        new_list.append(pop_min(num_list))
    return new_list
    ### END SOLUTION
print(simple_sort([6,8,3,9,2,0,3,0,1,2,1]))
```

<!-- grade_id: cell-6e005ec6d0f955c4, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

(c) If your function doesn't already do this, amend it so that as well as returning the sorted list, it has the side-effect of making the original list <em>equal</em> to the sorted list. So that if you type 

```python
this_list = [3, 5, 2, 7, 1, 2, 3, 1, 9]
simple_sort(this_list)
```

then you should get the output <code>[1, 1, 2, 2, 3, 3, 5, 7, 9]</code>, and as a side-effect, the value of <code>this_list</code> should become

<code>[1, 1, 2, 2, 3, 3, 5, 7, 9]</code>.

<!-- grade_id: cell-d8a06ebb274f04d6, points: 0, grade: false, solution: true -->
```python
def simple_sort(num_list):
    ### BEGIN SOLUTION
    # length of list
    n = len(num_list)
    # empty new list
    new_list = []
    # for loop
    for r in range(n):
        # pop the minimum and append it to new list
        new_list.append(pop_min(num_list))
    num_list += new_list
    return new_list
    ### END SOLUTION
my_list = [6,8,3,9,2,0,3,0,1,2,1]
print(simple_sort(my_list))
print(my_list)
```

<!-- grade_id: cell-b3f8d41096690235, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Adapt your <code>pop_min</code> function so that every time one element is compared with another, it increments by one a global variable called <code>comparison_count</code>. Hence if you type

```python
comparison_count = 0    
pm = pop_min([3, 5, 2, 7, 1, 2, 3, 1, 9])
comparison_count
```

you should get 9.

Using this, calculate the total number of comparisons necessary to sort the list <code>[3, 5, 2, 7, 1, 2, 3, 1, 9]</code> using <code>simple_sort</code>. 

Do you think this is an efficient method of sorting?

<!-- grade_id: cell-c05510e3a3d77323, points: 0, grade: false, solution: true -->
```python
def pop_min(num_list):
    ### BEGIN SOLUTION
    global comparison_count
    # initialize min at infinity
    this_min = num_list[0]
    min_index = 0
    # for loop
    for i in range(1,len(num_list)):
        # adjust minimum if necessary
        comparison_count+=1
        if num_list[i] < this_min:
            min_index = i
            this_min = num_list[i]
    return num_list.pop(min_index)
    ### END SOLUTION
```

<!-- grade_id: cell-ded140af0df44312, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
comparison_count = 0
question1d(_globals)
```

## Question 2

(i) Import the NumPy module, and then define three lists and three arrays as follows:

```python
int_list = [1, 2, 3]
float_list = [4.0, 5.0, 6.0]
int_float_list = [7, 8, 9.0]
int_list2 = [10, 11, 12]

int_arr = np.array(int_list)
float_arr = np.array(float_list)
int_float_arr = np.array(int_float_list)
int_arr2 = np.array(int_list2)
```

Print each of your four arrays, and extract elements of them using, for example
```python
int_arr[0]
```

Then:

(a) Complete the following statement:

A NumPy array based on a list of ints contains 

<ol>
    <li>only ints of some kind;</li>
    <li>only floats of some kind;</li>
    <li>both ints and floats.</li>
</ol>

<!-- grade_id: cell-251be60ac5fa1475, points: 0, grade: false, solution: true -->
```python
question2ia_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question2ia_answer = 1
### END SOLUTION
```

<!-- grade_id: cell-dc93660e40f37238, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2ia_answer == question2_ia())
print('test case passed!')
```

(b) Complete the following statement: 

A NumPy array based on a list of floats contains 

<ol>
    <li>only ints of some kind;</li>
    <li>only floats of some kind;</li>
    <li>both ints and floats.</li>
</ol>

<!-- grade_id: cell-81094386692ef9ea, points: 0, grade: false, solution: true -->
```python
question2ib_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question2ib_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-b5ad06d8b696439b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2ib_answer == question2_ib())
print('test case passed!')
```

(c) Complete the following statement: 

A NumPy array based on a list of ints and floats contains 

<ol>
    <li>only ints of some kind;</li>
    <li>only floats of some kind;</li>
    <li>both ints and floats.</li>
</ol>

<!-- grade_id: cell-6e960ccc047fbc69, points: 0, grade: false, solution: true -->
```python
question2ic_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question2ic_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-e2410de2feb52905, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2ic_answer == question2_ic())
print('test case passed!')
```

(d) Execute the following code:

```python
print(type(1))
print(type(int_arr[0]))
print(type(7.0))
print(type(int_float_arr[0]))
```

Which of the following statements is true?

<ol>
    <li>Default NumPy ints and floats are each of the same type as default native Python ints and floats.</li>
    <li>Default NumPy ints are of the same type as default native Python ints, but the same is not true of floats.</li>
    <li>Default NumPy floats are of the same type as default native Python floats, but the same is not true of ints.</li>
    <li>Default NumPy ints and floats are each of a different type from default native Python ints and floats.</li>.
</ol>

<!-- grade_id: cell-64db3c185956cedb, points: 0, grade: false, solution: true -->
```python
#question2id_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question2id_answer = 4
### END SOLUTION
```

<!-- grade_id: cell-8dbfcb744fca158e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2id_answer == question2id())
print('test case passed!')
```

(ii) Define the following arrays:

```python
int_float_complex_arr = np.array([1, 2.0, 3+0j])
int_float_complex_string_arr = np.array([1, 2.0, 3+0j, 'four'])
```

(a) Complete the following statement:

A NumPy array based on a list of ints, floats and complexes contains 

<ol>
    <li>only ints of some kind;</li>
    <li>only floats of some kind;</li>
    <li>only complexes of some kind;</li>
    <li>a mixture of ints, floats and complexes.</li>
</ol>

<!-- grade_id: cell-44f872ebef2c7162, points: 0, grade: false, solution: true -->
```python
question2iia_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question2iia_answer = 3
### END SOLUTION
```

<!-- grade_id: cell-77fb9047990820f7, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2iia_answer == question2_iia())
print('test case passed!')
```

(b) Complete the following statement:

A NumPy array based on a list of ints, floats, complexes and strings contains 

<ol>
    <li>only ints of some kind;</li>
    <li>only floats of some kind;</li>
    <li>only complexes of some kind;</li>
    <li>only strings of some kind;</li>
    <li>a mixture of ints, floats, complexes and strings.</li>
</ol>

<!-- grade_id: cell-3dd3bd52726607d0, points: 0, grade: false, solution: true -->
```python
question2iib_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question2iib_answer = 4
### END SOLUTION
```

<!-- grade_id: cell-60903a4c7897151d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2iib_answer == question2_iib())
print('test case passed!')
```

(c) Execute the following code:

```python
print(type(1+0j))
print(type(int_float_complex_arr[0]))
print(type('1'))
print(type(int_float_complex_string_arr[0]))
```

Which of the following statements is true?

<ol>
    <li>Default NumPy complexes and strings are each of the same type as default native Python complexes and strings.</li>
    <li>Default NumPy complexes are of the same type as default native Python complexes, but the same is not true of strings.</li>
    <li>Default NumPy strings are of the same type as default native Python strings, but the same is not true of complexes.</li>
    <li>Default NumPy complexes and strings are each of a different type from default native Python complexes and strings.</li>.
</ol>

<!-- grade_id: cell-19ee93a7adac4856, points: 0, grade: false, solution: true -->
```python
question2iic_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question2iic_answer = 4
### END SOLUTION
```

<!-- grade_id: cell-969a47dad48ec30a, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2iic_answer == question2_iic())
print('test case passed!')
```

## Question 3

(i) Calculate
```python
int_arr + float_arr
```
and similar "array sums". Then:

(a) Complete the following statement:

When you add two NumPy arrays together, you

<ol>
    <li>concatenate them (that is, join them head to tail, in the same way as lists or tuples);</li>
    <li>add them component-by-component, as if they were vectors.</li>
</ol>

<!-- grade_id: cell-a4802571a7742aba, points: 0, grade: false, solution: true -->
```python
question3ia_answer = # (delete, and fill in the value 1or 2)
### BEGIN SOLUTION
question3ia_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-1a957a83c20b93a8, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3ia_answer == question3_ia())
print('test case passed!')
```

(b) Complete the following statement:

The sum of an array of ints and an array of floats is

<ol>
    <li>an array of ints;</li>
    <li>an array of floats;</li>
    <li>an array containing both ints and floats.</li>
</ol>

<!-- grade_id: cell-ffe39e4933876bf7, points: 0, grade: false, solution: true -->
```python
question3ib_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3ib_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-309bc99958fb8e3b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3ib_answer == question3_ib())
print('test case passed!')
```

(c) Complete the following statement:

When you add two one-dimensional NumPy arrays of different lengths,

<ol>
    <li>it works fine;</li>
    <li>it generates an error message.</li>
</ol>

<!-- grade_id: cell-8cb815a797c62fab, points: 0, grade: false, solution: true -->
```python
question3ic_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3ic_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-d43f23b47e0b22e6, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3ic_answer == question3_ic())
print('test case passed!')
```

(ii) Experiment with using the other arithmetic operators, ```-```, ```*```, ```/```, ```//``` and ```**``` on scalars and 1D NumPy arrays. Then:

(a) Complete the following statement:

When you multiply a one-dimensional NumPy array by 3, you

<ol>
    <li>create three copies of the array, joined head to tail, as with lists or tuples;</li>
    <li>multiply each of the elements of the array separately by 3;</li>
    <li>multiply only the first element (that is, the element with index 0) by 3;</li>
    <li>none of the above: there's an error message.</li>
</ol>

<!-- grade_id: cell-b1f4300e476618ed, points: 0, grade: false, solution: true -->
```python
question3iia_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question3iia_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-237e0bbc53d9d43e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iia_answer == question3_iia())
print('test case passed!')
```

(b) In which of the following circumstances can you use Python's arithmetic operators?

I. To combine a scalar with a scalar.<br/>
II. To combine a scalar with a 1D NumPy array of any length.<br/>
III. To combine two 1D NumPy arrays of the same length.<br/>
IV. To combine two 1D NumPy arrays of different lengths.

<ol>
    <li>All of them.</li>
    <li>All of them except II.</li>
    <li>All of them except IV.</li>
    <li>All of them except II and IV.</li>
</ol>

<!-- grade_id: cell-270dbaddd730484f, points: 0, grade: false, solution: true -->
```python
question3iib_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question3iib_answer = 3
### END SOLUTION
```

<!-- grade_id: cell-77c15fb7ccfeb3d6, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iib_answer == question3_iib())
print('test case passed!')
```

(iii) Create the 2D NumPy arrays

```python
mat33 = np.array([[1,2,-1],[3,-4,0],[2,-1,1]])
mat13 = np.array([[1,2,3]])
mat31 = np.array([[1],[2],[3]])
mat23 = np.array([[1,3,5],[1,2,4]])
```

Experiment, and consider the following claims:

(a) When a $3\times 3$ NumPy array is added to a $1 \times 3$ NumPy array:

<ol>
    <li>an error is generated;</li>
    <li>it works, and the $1 \times 3$ array is treated as if it were three repeated rows;</li>
    <li>it works, and the $1 \times 3$ array is treated as if it were three repeated columns.</li>
</ol>

<!-- grade_id: cell-a0594c26e4e367f1, points: 0, grade: false, solution: true -->
```python
question3iiia_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3iiia_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-b4cf36b001d570fc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iiia_answer == question3_iiia())
print('test case passed!')
```

(b) When a $3\times 3$ NumPy array is added to a $3 \times 1$ NumPy array:

<ol>
    <li>an error is generated;</li>
    <li>it works, and the $3 \times 1$ array is treated as if it were three repeated rows;</li>
    <li>it works, and the $3 \times 1$ array is treated as if it were three repeated columns.</li>
</ol>

<!-- grade_id: cell-d6767f4dcf44ca69, points: 0, grade: false, solution: true -->
```python
question3iiib_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3iiib_answer = 3
### END SOLUTION
```

<!-- grade_id: cell-47c7fa2c91873e9d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iiib_answer == question3_iiib())
print('test case passed!')
```

(c) When a $3\times 3$ NumPy array is added to a $2 \times 3$ NumPy array:

<ol>
    <li>an error is generated;</li>
    <li>it works, with rows being repeated cyclically as necessary;</li>
    <li>it works, with columns being repeated cyclically as necessary;</li>
</ol>

<!-- grade_id: cell-f3e24333aa57f4e8, points: 0, grade: false, solution: true -->
```python
question3iiic_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3iiic_answer = 1
### END SOLUTION
```

<!-- grade_id: cell-589aa4ba9057b080, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iiic_answer == question3_iiic())
print('test case passed!')
```

(d) When a $3\times 3$ NumPy array is added to a 1D array of length 3:

<ol>
    <li>an error is generated;</li>
    <li>it works, and you get the same result as if the 1D array were a $1 \times 3$ array;</li>
    <li>it works, and you get the same result as if the 1D array were a $3 \times 1$ array.</li>
</ol>

<!-- grade_id: cell-6d3c2880efdacf91, points: 0, grade: false, solution: true -->
```python
question3iiid_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3iiid_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-5ef0a020bd7e8f5d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iiid_answer == question3_iiid())
print('test case passed!')
```

(iv) Experiment with commands like
```python
mat33 == mat33
mat31 < mat33
mat33 >= int_arr
mat23 > mat33
```
and so on.

Then:

(a) Comparing two identical NumPy arrays using ```==```:

<ol>
    <li>creates the output <code>True</code>, just as if we were comparing lists or tuples;</li>
    <li>creates an array with the same number of rows and columns, every entry of which is <code>True</code>;</li>
    <li>generates an error message.</li>
</ol>

<!-- grade_id: cell-7ad8b0120c388fca, points: 0, grade: false, solution: true -->
```python
question3iva_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question3iva_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-48bfbbf9e696025f, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3iva_answer == question3_iva())
print('test case passed!')
```

(b) True or False: two numerical NumPy arrays can be compared using ```==```, ```<```, ```>```, ```<=``` or ```>=``` if and only if they can be combined using ```+```, ```-```, ```*```, ```/```, ```//``` or ```**```.

<!-- grade_id: cell-da472d31bb8ebb72, points: 0, grade: false, solution: true -->
```python
question3ivb_answer = # (delete, and fill in the value True or False)
### BEGIN SOLUTION
question3ivb_answer = True
### END SOLUTION
```

<!-- grade_id: cell-bacfe302970cbe90, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question3ivb_answer == question3_ivb())
print('test case passed!')
```

## Question 4

(a) Complete the following statement. Once NumPy has been imported as `np`, the values of
```python
np.linspace(7, 101, 95)
```
and
```python
np.arange(7, 102, 1)
```
are

<ol>
    <li>identical in every respect;</li>
    <li>equal (in all components), but not quite identical;</li>
    <li>neither equal nor identical.</li>
</ol>

<!-- grade_id: cell-62e2d095683626b7, points: 0, grade: false, solution: true -->
```python
question4a_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question4a_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-57aa3854b3070f04, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4a_answer == question4_a())
print('test case passed!')
```

(b) Complete the following statement. Once NumPy has been imported as `np`, the values of
```python
np.sqrt(np.linspace(7, 101, 95))
```
and
```python
np.sqrt(np.arange(7, 102, 1))
```
are

<ol>
    <li>identical in every respect;</li>
    <li>equal (in all components), but not quite identical;</li>
    <li>neither equal nor identical.</li>
</ol>

<!-- grade_id: cell-c97303140529b98e, points: 0, grade: false, solution: true -->
```python
question4b_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question4b_answer = 1
### END SOLUTION
```

<!-- grade_id: cell-d8ec207d1d90bfc5, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4b_answer == question4_b())
print('test case passed!')
```

(c) Using either `linspace` or `arange`, create a 1D NumPy array called `theta_array` consisting of 401 equally-spaced values between $0$ and $2\,\pi$.

Then create another 1D NumPy array called `x_array` consisting of $\cos(3\,\theta)$ for each $\theta$ in `theta_array` (you shouldn't need to use a loop or a comprehension for this, though autograding won't be able to tell if you have).

Create, too, a 1D NumPy array called `y_array` consisting of $\sin(5\,\theta)$ for each $\theta$ in `theta_array`.

Finally, create a plot of `x_array` (on the horizontal axis) against `y_array` (on the vertical axis).

<!-- grade_id: cell-af56d68000c9fd5f, points: 0, grade: false, solution: true -->
```python
%matplotlib inline
import matplotlib.pyplot as plt

### BEGIN SOLUTION
theta_array = np.linspace(0,2*np.pi,401)
x_array = np.cos(3*theta_array)
y_array = np.sin(5*theta_array)
fig, ax = plt.subplots()
plt.plot(x_array,y_array)
### END SOLUTION
```

<!-- grade_id: cell-df35ac4da0a1a0fd, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4c(_globals)
```

<!-- grade_id: cell-6d1cec33ed26d6cc, points: 1, grade: true, solution: false -->
```python
# 1 Mark
_globals = globals()
question4c_plot(_globals)
```

## Question 5

(i) Investigate expressions like
```python
np.dot(mat23, mat33)
np.dot(mat33, mat31)
np.dot(mat33, int_arr)
np.dot(int_arr, int_arr2)
```
etc. Then:

(a) The `dot` function in the NumPy module can be used to find

<ol>
    <li>the scalar product of two vectors represented as 1D arrays;</li>
    <li>the matrix product of two matrices represented as 2D arrays;</li>
    <li>both.</li>
</ol>

<!-- grade_id: cell-b08d42f459524526, points: 0, grade: false, solution: true -->
```python
question5ia_answer = # (delete, and fill in the value 1, 2 or 3)
### BEGIN SOLUTION
question5ia_answer = 3
### END SOLUTION
```

<!-- grade_id: cell-0c393b69be5aaeaf, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5ia_answer == question5_ia())
print('test case passed!')
```

(b) Which of the following work; that is, don't generate an error message? 

<ol>
    <li><code>np.dot(mata, matb)</code>, where <code>mata</code> and <code>matb</code> are $3\times 3$ matrices?</li>
    <li><code>np.dot(mata, matb)</code>, where <code>mata</code> is a $3 \times 3$ matrix and <code>matb</code> is a $3 \times 1$ matrix?</li>
    <li><code>np.dot(mata, matb)</code>, where <code>mata</code> is a $3 \times 3$ matrix and <code>matb</code> is a $1 \times 3$ matrix?</li>
    <li><code>np.dot(mata, matb)</code>, where <code>mata</code> is a $1 \times 3$ matrix and <code>matb</code> is a $3 \times 3$ matrix?</li>
    <li><code>np.dot(mata, matb)</code>, where <code>mata</code> is a $3 \times 1$ matrix and <code>matb</code> is a $3 \times 3$ matrix?</li>
</ol>

Type in your answers in curly brackets separated by commas; for example, if you think 1, 3 and 5 are the correct answers, type in `{1, 3, 5}`.

<!-- grade_id: cell-4b8f3784cbe2da4c, points: 0, grade: false, solution: true -->
```python
question5ib_answer = # (delete, and fill in your values)
### BEGIN SOLUTION
question5ib_answer = {1, 2, 4}
### END SOLUTION
```

<!-- grade_id: cell-9f915f532f6cf07c, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5ib_answer == question5_ib())
print('test case passed!')
```

(c) Suppose `mat` is a $3 \times 3$ NumPy array and `vec` is a 1D array of length 3; numerical in each case. Then the expressions

```python
np.dot(mat, vec)
np.dot(vec, mat)
```

<ol>
    <li>both work, and generate the same output;</li>
    <li>both work, and generate different output;</li>
    <li>work in the first case, though the second generates an error message;</li>
    <li>work in the second case, though the first generates an error message;</li>
    <li>both fail, generating error messages.</li>

<!-- grade_id: cell-5be9b1a83a30a998, points: 0, grade: false, solution: true -->
```python
question5ic_answer = # (delete, and fill in the value 1, 2, 3, 4 or 5)
### BEGIN SOLUTION
question5ic_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-6980e868995ee762, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5ic_answer == question5_ic())
print('test case passed!')
```

(d) Suppose `mat` is a $3 \times 3$ NumPy array and `vec` is a 1D array of length 3; numerical in each case. Then, in the expressions

```python
np.dot(mat, vec)
np.dot(vec, mat)
```

the 1D array `vec` is interpreted by Python

<ol>
    <li>as a row vector in both cases;</li>
    <li>as a column vector in both cases;</li>
    <li>as a row vector in the first case and a column vector in the second case;</li>
    <li>as a column vector in the first case and a row vector in the second case.</li> 
</ol>

<!-- grade_id: cell-e4ba1626c13d76c0, points: 0, grade: false, solution: true -->
```python
question5id_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question5id_answer = 4
### END SOLUTION
```

<!-- grade_id: cell-b4fae0f7029c8219, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question5id_answer == question5_id())
print('test case passed!')
```

(ii) The matrix $A$ is defined by

$$A = \left(\begin{array}{cccc}
1&-3&2&2\\0&1&0&-1\\1&3&1&-2\\-2&0&-3&-2
\end{array}\right)$$

and the vector ${\bf v}$ by

$${\bf v} = \left(\begin{array}{c}
0\\-2\\-5\\-1
\end{array}\right).$$

(a) Set them up in Python as, respectively, a $4 \times 4$ NumPy array and a 1D NumPy array of length 4.

<!-- grade_id: cell-c73166e37cc04512, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
A = np.array([[1,-3,2,2],[0,1,0,-1],[1,3,1,-2],[-2,0,-3,-2]])
v = np.array([0, -2, -5, -1])
### END SOLUTION
```

<!-- grade_id: cell-4abdd3e56000e898, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iia(_globals)
```

Import the `linalg` submodule of the NumPy module, and using functions from it, calculate:

(b) The norm, or magnitude, of the vector ${\bf v}$, $|{\bf v}|$ (sometimes written $||{\bf v}||$). Call this `normv`.

<!-- grade_id: cell-91f9867623ca4977, points: 0, grade: false, solution: true -->
```python
import numpy.linalg as la
### BEGIN SOLUTION
normv = la.norm(v)
### END SOLUTION
print(normv)
```

<!-- grade_id: cell-ee2f35e08eee306a, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iib(_globals)
```

(c) The 1-norm of ${\bf v}$, $||{\bf v}||_1$, equal to the sum of the absolute values of ${\bf v}$'s components. Call this `norm1v`.

<!-- grade_id: cell-f13390acb597945e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
norm1v = la.norm(v,1)
### END SOLUTION
print(norm1v)
```

<!-- grade_id: cell-debc459b8c1d3bfd, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iic(_globals)
```

(d) The $\infty$-norm of ${\bf v}$, $||{\bf v}||_\infty$, equal to the maximum absolute value of ${\bf v}$'s components. Call this `norminfv`.

<!-- grade_id: cell-93aadf5cb000e879, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
norminfv = la.norm(v,np.Inf)
### END SOLUTION
print(norminfv)
```

<!-- grade_id: cell-19a9351daa840d32, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iid(_globals)
```

(e) The determinant of $A$. Call this `detA`.

<!-- grade_id: cell-937de8cbc2599027, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
detA = la.det(A)
### END SOLUTION
print(detA)
```

<!-- grade_id: cell-569be2cbbb9eed1d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iie(_globals)
```

(f) The inverse matrix of $A$. Call this `invA`.

<!-- grade_id: cell-7219a649ec63e886, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
invA = la.inv(A)
### END SOLUTION
print(invA)
```

<!-- grade_id: cell-90c1abd85cc154bd, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iif(_globals)
```

(g) The eigenvalues of A. Call these `eigA`.

<!-- grade_id: cell-dab55830ccdc8281, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
eigA = la.eigvals(A)
### END SOLUTION
print(eigA)
```

<!-- grade_id: cell-19e5bfce7227c277, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iig(_globals)
```

(h) The value of $A^5$, where the power indicates repeated <em>matrix</em> multiplication. (You can do this using repeated applications of `dot`, and autograding won't be able to tell if you have, but try to do it using a function from the `linalg` submodule instead.) Call this `Apow5`.

<!-- grade_id: cell-b110c338a988e323, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
Apow5 = la.matrix_power(A,5)
### END SOLUTION
print(Apow5)
```

<!-- grade_id: cell-1f8920c55649c114, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iih(_globals)
```

(j) (and I hope you can see why I've skipped part (i)...) The solution, ${\bf x}$, of the linear equation system

$$A\,{\bf x} = {\bf v},$$

as a 1D NumPy array. Call this `solution1`.

<!-- grade_id: cell-6fa6caf023f1ae4c, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
solution1 = la.solve(A,v)
### END SOLUTION
print(solution1)
```

<!-- grade_id: cell-5f443bea88b9db25, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iij(_globals)
```

(k) The column vector ${\bf w}$ is defined by

$${\bf w} = \left(\begin{array}{cccc}
1&0&1&-1
\end{array}\right).$$

Find, as a 1D NumPy array, the solution ${\bf y}$ of the linear equation system

$${\bf y} \, A = {\bf w}$$

(this may take some thought). Call this `solution2`

<!-- grade_id: cell-48f16eda5eeeea84, points: 0, grade: false, solution: true -->
```python
w =np.array([1,0,1,-1])
### BEGIN SOLUTION
solution2 = la.solve(np.transpose(A),w)
### END SOLUTION
print(solution2)
```

<!-- grade_id: cell-3fa9ca4d1f5124a8, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question5iik(_globals)
```

## Question 6

(a) Define a `poly1d` object called `mypoly` corresponding to the quartic 
$$x^4-8\,x^3-26\,x^2+168\,x-135.$$

<!-- grade_id: cell-04b63033fa05f1b4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mypoly = np.poly1d([1,-8,-26,168,-135])
### END SOLUTION
print(mypoly)
```

<!-- grade_id: cell-1b66e2e68a6b981d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question6a(_globals)
```

(b) Calculate the values of `mypoly` at $x=-6,-5,-4,\dots,8,9,10$, giving your answer in the form of a NumPy array called `mypolyvalues`.

<!-- grade_id: cell-1b241ac1c6841df4, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mypolyvalues = np.polyval(mypoly,np.linspace(-6,10,17))
### END SOLUTION
print(mypolyvalues)
```

<!-- grade_id: cell-21da96f121ceef32, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question6b(_globals)
```

(c) Calculate the roots of `mypoly`, giving your answer in the form of a NumPy array called `mypolyroots`.

<!-- grade_id: cell-eddd0657cedbaf6d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mypolyroots = np.roots(mypoly)
### END SOLUTION
print(mypolyroots)
```

<!-- grade_id: cell-7cdca97652b4de07, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question6c(_globals)
```

(d) Calculate the $x$- and $y$-coordinates of the stationary points of `mypoly`, giving your answers as NumPy arrays `spx` and `spy`.

<!-- grade_id: cell-a7f7dfea33663a11, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
spx = np.roots(np.polyder(mypoly))
spy = np.polyval(mypoly,spx)
### END SOLUTION
print(spx)
print(spy)
```

<!-- grade_id: cell-785811b5247a03e3, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question6d(_globals)
```

(e) Create a plot of `mypoly` for $-6\le x \le 10$, showing the positions of the three stationary points you found above. Use exactly 200 points on the plot.

<!-- grade_id: cell-86800ef188199ebe, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
xvals = np.linspace(-6,10,200)
yvals = np.polyval(mypoly, xvals)
fig, ax = plt.subplots()
plt.plot(xvals,yvals,'',spx,spy,'.r')
### END SOLUTION
from plotchecker import LinePlotChecker
pc = LinePlotChecker(ax)
len(pc.x_data[1])
```

<!-- grade_id: cell-2a1ac9e55709699f, points: 2, grade: true, solution: false -->
```python
# 2 Marks
_globals = globals()
question6e_plot(_globals)
```

## Bonus Question, Part 1

Write a function `root_finder`, which takes as inputs the lambda-expressions `f` and `df` (representing a function and its derivative), the left and right endpoints `a` and `b` respectively, and a tolerance `tol`. The function should implement BOTH the bisection method and Newton's method; the bisection method should use the initial interval `(a, b)`, and Newton's method should use the initial value `(a + b)/2`. 

In the case of the bisection method, iteration should continue for as long as the interval length is greater than the tolerance; in the case of Newton's method, iteration should continue for as long as the difference between successive iterates is greater than the tolerance.

Your function should return the final approximation to the root as well as the number of iterations required to achieve the approximation for each method. The returned values should be of the form: `bisection_approx, bisection_iterations, newton_approx, newton_iterations`. The bisection search should return the midpoint of the final interval.

<!-- grade_id: cell-82740cb38e52ea64, points: 0, grade: false, solution: true -->
```python
def root_finder(f, g, a, b, tol):
    ### BEGIN SOLUTION
    return (bisection(f, a, b, tol) + newton_method(f, g, (a+b)/2, tol))      
    
def bisection(f, a, b, tol):
        it = 0
        while b-a > tol:
                c = 0.5*(a + b)
                if f(a)*f(c) < 0:
                    b = c
                else:
                    a = c
                it += 1
        return 0.5*(a + b), it
    
    
def newton_method(f, g, x0, tol):
        xnew = x0 - (f(x0)/g(x0))
        it   = 1
        while abs(xnew - x0) > tol:
            x0   = xnew
            xnew = x0 - (f(x0)/g(x0))
            it  += 1
        return xnew, it
     ### END SOLUTION

```

<!-- grade_id: cell-b261b86d49ae6b3c, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
f1     = lambda x: x**2 - 2
df1     = lambda x: 2*x
a1     = 0
b1     = 3
tol1   = 1e-7

npt.assert_almost_equal(1.4142135232686996, root_finder(f1, df1, a1, b1, tol1)[0])
print("Test case: correct estimate returned by bisection method!")
```

<!-- grade_id: cell-c8cce025f525fa72, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
f1     = lambda x: x**2 - 2
df1     = lambda x: 2*x
a1     = 0
b1     = 3
tol1   = 1e-7

bisection_its = root_finder(f1, df1, a1, b1, tol1)[1]
assert(bisection_its >= 24 and bisection_its <= 26)
print("Test case: right number of iterations performed by bisection method!")
```

<!-- grade_id: cell-1c15c3b7067120e0, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
f1     = lambda x: x**2 - 2
df1     = lambda x: 2*x
a1     = 0
b1     = 3
tol1   = 1e-7

npt.assert_almost_equal(1.4142135623730951, root_finder(f1, df1, a1, b1, tol1)[2])
print("Test case: correct estimate returned by Newton's method!")
```

<!-- grade_id: cell-80848991afe48990, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 Mark
f1     = lambda x: x**2 - 2
df1     = lambda x: 2*x
a1     = 0
b1     = 3
tol1   = 1e-7

newton_its = root_finder(f1, df1, a1, b1, tol1)[3]
assert(newton_its >= 3 and newton_its <= 5)
print("Test case: right number of iterations performed by Newton's method!")
```

<!-- grade_id: cell-30cbf63266ef19bd, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 Marks
### BEGIN HIDDEN TESTS
f2 = lambda x: np.exp(x) - 2*x**2 + 3*x
df2 = lambda x: 3 + np.exp(x) - 4*x
a2 = -2
b2 = 1
tol2 = 1e-7

rf = root_finder(f2, df2, a2, b2, tol2)
assert(
    np.isclose(-0.22973136603832245, rf[0]) and rf[1]>=24 and rf[1]<=26 and np.isclose(-0.22973135981918785, rf[2]) and rf[3] >= 3 and rf[3] <= 5
)
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

## Bonus Question, Part 2

Usually, Newton's method outperforms interval bisection, but this isn't invariably the case. Investigate the convergence of both methods to the zero, between $x=0$ and $x=1$, of $f_2(x)=\sin^3(x+3)$; use $a=0$ and $b=1$, and a tolerance of $10^{-7}$.

Why exactly does Newton's method do so badly here?

Using what you've learnt, find a polynomial function $f_3$ with a zero at exactly $x=\sqrt{2}$ to which both methods converge (using $a=1$, $b=2$ and a tolerance of $10^{-7}$) but to which Newton's method converges <em>much</em> more slowly than interval bisection, taking more than twice as many iterations.

<!-- grade_id: cell-3ebd1f2ce560178f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
f3 = lambda x: (x**2-2)**5
df3 = lambda x: 10*x*(x**2-2)**4
### END SOLUTION
```

<!-- grade_id: cell-9c964721a4dd68a7, points: 2, grade: true, solution: false -->
```python
a3 = 1
b3 = 2
tol3 = 1e-7

rf3 = root_finder(f3, df3, a3, b3, tol3)
assert(rf3[3] > 2*rf3[1])
print("Test passed!")
```

```python

```