```python
from autograder_term2week2 import *
_globals = globals()
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki</center>

# <center> Worksheet 6: More About Functions</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-568a0b8879ff2f43, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

## Question 1

For a planet or comet, the <b>mean anomaly</b>, $M$, is the time since the closest approach to the Sun, divided by the length of the year, multiplied by $2\pi$ (so for example if one-quarter of the year has gone by, $M=\pi/2$). The <b>eccentric anomaly</b>, $E$, is a quantity we use to locate the planet or comet on its orbit. 
    
<b>Kepler's Equation</b> links the two. If the <b>orbital eccentricity</b> is $e$, then Kepler's Equation is
$$M=E-e\,\sin E.$$

(a) For a comet, the eccentricity $e$ can be quite large. Setting $e=0.95$ and $M=\pi/2$, estimate the value of $E$ by iterating the function
$$M+e\,\sin E$$
10 times, starting from the value $E=M$, and <em>making sure to use a lambda-expression</em>.

The function `nestlist` from the notes will be useful, and is supplied. Call your list of 11 iterates `Elist1`.

<!-- grade_id: cell-f80e92e0900ca624, points: 0, grade: false, solution: true -->
```python
from math import pi, sin, cos

def nestlist(f, x0, n):
    "Returns iterates of x = f(x), x0, x1, ... xn"
    
    # initialise variable and list
    x = x0
    xlist = [x0]
    
    # for loop
    for r in range(n):
        # update variable
        x = f(x)
        # append latest value to list
        xlist.append(x)
    
    # return list
    return xlist

M, e = pi/2, 0.95
Elist1 = nestlist(lambda E: ...  # fill in your code here
### BEGIN SOLUTION
Elist1 = nestlist(lambda E: M + e*sin(E), M, 10)
### END SOLUTION
print(Elist1)
```

<!-- grade_id: cell-53518b4701ec028d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

(b) Estimate the mean anomaly again, but this time by instead iterating
$$\frac{M+e\,\sin E + E}{2}$$
10 times. Use the same eccentricity, $e=0.95$, the same mean anomaly, $M=\pi/2$ and the same starting value, $E=M$. Once more, use a lambda-expression in your call to `nestlist`. This time, call the list of 11 iterates `Elist2`.

<!-- grade_id: cell-f976fa3495d05ab2, points: 0, grade: false, solution: true -->
```python
M, e = pi/2, 0.95
### BEGIN SOLUTION
Elist2 = nestlist(lambda E: (M + e*sin(E) + E)/2, M, 10)
### END SOLUTION
print(Elist2)
```

<!-- grade_id: cell-6b2f0edf42d38cab, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

(c) Estimate the mean anomaly a third time, but this time by instead iterating
$$\frac{M+e\,\sin E - E\,e\,\cos E}{1- e\,\cos E}$$
10 times. Use the same eccentricity, $e=0.95$, the same mean anomaly, $M=\pi/2$ and the same starting value, $E=M$. Once more, use a lambda-expression in your call to `nestlist`. This time, call the list of 11 iterates `Elist3`.

<!-- grade_id: cell-014d16016afd87ab, points: 0, grade: false, solution: true -->
```python
M, e = pi/2, 0.95
### BEGIN SOLUTION
Elist3 = nestlist(lambda E: (M + e*sin(E) - E*e*cos(E))/(1-e*cos(E)), M, 10)
### END SOLUTION
print(Elist3)
```

<!-- grade_id: cell-f304b70d40d75159, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Write and test a function called `eccentric_anomaly`, which takes as its arguments floats $M$ and $e$, and a non-negative integer $M$, and returns, as a float, an estimate for the eccentric anomaly associated with mean anomaly $M$ and eccentricity $e$. This should be obtained from Kepler's equation by iterating $n$ times, starting with $E=M$, using the most rapidly converging of the three iterations above. You may assume that $0\le e < 1$.

<!-- grade_id: cell-0edaadf7edf15359, points: 0, grade: false, solution: true -->
```python
def eccentric_anomaly(M, e, n):
    from math import pi, sin, cos
    """
    Returns an estimate for the eccentric anomaly associated
    with mean anomaly M, and eccentricity e, obtained by n
    iterations of a rapidly converging sequence
    """
    ### BEGIN SOLUTION
    # initialize
    E = M
    
    # for loop
    for r in range(n):
        E = (M + e*sin(E) - E*e*cos(E))/(1-e*cos(E))
    
    # return final value
    return E
    ### END SOLUTION
```

<!-- grade_id: cell-4c4e58d7e8c3d46e, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

## Question 2

(a) Write a function called <code>integer_digits</code>, which takes as its arguments <code>number</code> (an int greater than or equal to 0) and <code>base</code> (an int greater than 1), and returns, as a list, the digits of <code>number</code> in base <code>base</code>. 

Test your function: the commands <code>integer_digits(347, 10)</code> and <code>integer_digits(347, 2)</code> should return, respectively, <code>[3, 4, 7]</code> and <code>[1, 0, 1, 0, 1, 1, 0, 1, 1]</code>. If <code>number</code> is equal to zero, an empty list should be returned.

Make sure your code also works if you call it using <b>keywords</b>: for example,

```python
integer_digits(number = 347, base = 10)
```

<!-- grade_id: cell-6c432b941e2929af, points: 0, grade: false, solution: true -->
```python
def integer_digits(number, base):
    ### BEGIN SOLUTION
    # initialize list of digits
    digit_list = []
    # while loop
    while number > 0:
        # insert last digit
        digit_list.insert(0, number % base)
        # throw away last digit from number
        number //= base
    return digit_list
    ### END SOLUTION
```

<!-- grade_id: cell-f96bf7f1f098f4a6, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(integer_digits(347,2) == [1, 0, 1, 0, 1, 1, 0, 1, 1])
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-5922f55d871755cd, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

(b) Adapt your <code>integer_digits</code> function so that the default value of <code>base</code> is 10. 

Test. If you've done it right, 

```python
integer_digits(347, 10)
integer_digits(347, 2)
integer_digits(number = 347, base = 10)
```

and

```python
integer_digits(base = 10, number = 347)
```
should all still work, but in addition, 

```python
integer_digits(347)
```

and 

```python
integer_digits(number = 347)
```

should both return <code>[3, 4, 7]</code>.

<!-- grade_id: cell-b2b38ba00344d53d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
def integer_digits(number, base=10):
    # initialize list of digits
    digit_list = []
    # while loop
    while number > 0:
        # insert last digit
        digit_list.insert(0, number % base)
        # throw away last digit from number
        number //= base
    return digit_list
### END SOLUTION
```

<!-- grade_id: cell-a503a65776131ce2, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(integer_digits(347) == [3,4,7])
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-3a23b2de29ebc595, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2b(_globals)
```

(c) Adapt your code further so that <code>base</code> is only accepted as a keyword argument, but <code>number</code> can be positional or keyword.

If you've done it right, the following should work:

```python
integer_digits(347)
integer_digits(number = 347)
integer_digits(347, base = 2)
integer_digits(number = 347, base = 2)
integer_digits(base = 2, number = 347)
```

but 

```python
integer_digits(347, 2)
```

should no longer work.

Can you think of an advantage of this arrangement?

<!-- grade_id: cell-fda6ae637f58982c, points: 0, grade: false, solution: true -->
```python
## BEGIN SOLUTION
def integer_digits(number, *, base=10):
    # initialize list of digits
    digit_list = []
    # while loop
    while number > 0:
        # insert last digit
        digit_list.insert(0, number % base)
        # throw away last digit from number
        number //= base
    return digit_list
## END SOLUTION
```

<!-- grade_id: cell-70f2ba6cc7659b0a, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

## Question 3
(a) Write and test a hailstone sequence function which uses the infinite loop

```python
    while True:
```

instead of a loop with a terminating condition. This function should terminate at the first occurrence of 1. 

An <code>hailstone_step</code> function is provided for you below.

Note: autograding won't be able to tell whether you've used <code>while True</code>, but please do try.

<!-- grade_id: cell-74b97b68f8c335a0, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    if a % 2 == 0:
        return a//2
    else:
        return 3*a+1

def hailstone_sequence(a):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    if a==1:
        return [1]
    while True:
        # update and append
        a = hailstone_step(a)
        alist.append(a)
        # if sequence should now terminate, return alist and stop
        if a==1:
            return alist
    ### END SOLUTION
```

<!-- grade_id: cell-1f24b08e62f943b4, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

(b) Write and test an <code>aliquot_chain</code> function, called <code>aliquot_chain</code> which uses the infinite loop

```python
    while True:
```

instead of a loop with a terminating condition. It should terminate at the first occurrence of a number that has already been seen. You may want to reuse your <code>factor_sum</code> function from earlier in the module.

<!-- grade_id: cell-3d6fc02febf86f32, points: 0, grade: false, solution: true -->
```python
def factor_sum(a):
    ### BEGIN SOLUTION
    from math import floor, sqrt
    # initalize fsum
    fsum = 0
    # square root of a
    sqrta = floor(sqrt(a))
    # for loop
    for r in range(1,sqrta+1):
        if a%r == 0:
            fsum += r + a//r
    # adjust for perfect squares
    if sqrta ** 2 == a:
        fsum -= sqrta
    return fsum
    ### END SOLUTION

def aliquot_chain(a):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    while True:
        # update and append
        a = factor_sum(a)-a
        alist.append(a)
        # if number has already occurred, return alist and stop
        if a in alist[:-1]:
            return alist
    ### END SOLUTION
```

<!-- grade_id: cell-547043decd0b74d6, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

(c) The following code is for a function called <code>power_of_two_divisor1</code>, which takes as its argument <code>n</code>, assumed to be a positive integer, and returns the highest value of <code>r</code> such that $2^r$ divides $n$; so for example

```python
power_of_two_divisor1(24)
```

returns 3, because $2^3$ divides 24.

```python
def power_of_two_divisor1(n):
    # initialize r
    r = 0
    
    # while loop, terminating when n becomes odd
    while n % 2 == 0:
        # increment n
        r += 1
        # divide n by 2
        n = n // 2
    
    # return final value of r
    return r
```

Test this function.

Then, write and test a version of it called <code>power_of_two_divisor2</code> that doesn't use a while loop at all, but instead uses the infinite loop

```python
for r in count()
```

An outline is provided below.

<!-- grade_id: cell-7ca2fe4de8b88de3, points: 0, grade: false, solution: true -->
```python
def power_of_two_divisor2(n):
    from itertools import count
    # infinite for loop
    for r in count():
        ### BEGIN SOLUTION
        # if n is odd return latest value of r
        if not(n % 2 == 0):
            return r
            
        # otherwise, divide n by 2
        else:
            n = n // 2
        ### END SOLUTION
```

<!-- grade_id: cell-b763bf1de9d34a56, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3c(_globals)
```

## Question 4

(a) Complete the following factorial function, <code>my_factorial</code>, which works using <b>recursion</b>: that is, by calling itself. Test your function.

<!-- grade_id: cell-8e24c72afd376c6e, points: 0, grade: false, solution: true -->
```python
def my_factorial(n):
    # base case
    if n==1 or n==0:
        return 1
    # recursion step
    else:
        ### BEGIN SOLUTION
        return n * my_factorial(n - 1)
        ### END SOLUTION
```

<!-- grade_id: cell-d8620df67910e0fb, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

(b) Write and test a recursive version of your <code>hailstone_sequence</code> function, called <code>hailstone_sequence2</code>, which should work by calling itself instead of by a loop, and which should terminate at the first occurence of $1$. 

Autograding can't tell whether you've used recursion, but please do try!

<!-- grade_id: cell-483675ae72e172f2, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    if a % 2 == 0:
        return a//2
    else:
        return 3*a+1

def hailstone_sequence2(a):
    ### BEGIN SOLUTION
    # base case
    if a==1:
        return [1]
    # recursion step
    else:
        return [a] + hailstone_sequence2(hailstone_step(a))
    ### END SOLUTION
```

<!-- grade_id: cell-5b9dff1d06c22cb3, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4b(_globals)
```

(c) The following code defines an iterative function called <code>fibonacci</code>:

```python
def fibonacci(n):
    # initialize a and b
    a, b = 1, 1
    # the cases 1 and 2
    if n==1 or n==2:
        return 1
    else:
        # for loop
        for r in range(3,n+1):
            a, b = b, a + b
        return b
```

The following is an alternative, recursive version:

```python
def fibonacci2(n):
    # base case
    if n==1 or n==2:
        return 1
    # recursion step
    else:
        return fibonacci2(n-1) + fibonacci2(n-2)
```

Test that these two functions return the same arguments for $n=1, 2, 3, \dots, 10$.

Autograding can't tell whether you've done that or not, but please do!

Calculate the 35th Fibonacci number using each of these two functions, then choose one of the following options:

1. Both functions give the same output for $n=35$, in about the same time.
2. Both functions give the same output for $n=35$, but <code>fibonacci</code> is the more efficient.
3. Both functions give the same output for $n=35$, but <code>fibonacci2</code> is the more efficient.
4. The two functions give different outputs for $n=35$, suggesting at least one of them has a bug.

<!-- grade_id: cell-fe4c35facb2ba712, points: 0, grade: false, solution: true -->
```python
question4c_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question4c_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-1b5a4c306b8824ef, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4c_answer == question4_c())
print('test case passed!')
```

(d) Write and test a recursive function called <code>fibonacci_recpair</code> which takes as its argument <code>n</code>, assumed to be an integer greater than 1, and returns, as a 2-tuple, the $(n-1)$th and $n$th Fibonacci numbers. So, for example, 

```python
fibonacci_recpair(10)
```

should return <code>(34, 55)</code>.

<!-- grade_id: cell-4e6a39f259575fb1, points: 0, grade: false, solution: true -->
```python
def fibonacci_recpair(n):
    ### BEGIN SOLUTION
    # base case
    if n==2:
        return (1, 1)
    # recursion step
    else:
        previous = fibonacci_recpair(n-1)
        return (previous[1], previous[0]+previous[1])
    ### END SOLUTION
```

<!-- grade_id: cell-dc17b901d6f3b7f6, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4d(_globals)
```

(e) Write and test a function called <code>fibonacci3</code> which, by calling <code>fibonacci_recpair</code>, calculates the $n$th Fibonacci number. (This won't work for $n=1$, so you'll need to make a special case of that.)

<!-- grade_id: cell-2b18346684193918, points: 0, grade: false, solution: true -->
```python
def fibonacci3(n):
    ### BEGIN SOLUTION
    if n==1:
        return 1
    else:
        return fibonacci_recpair(n)[1]
    ### END SOLUTION
```

<!-- grade_id: cell-e258528bc9ebe868, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4e(_globals)
```

(f) Test your code from part (e) on $n=35$, and perhaps on higher values, and then choose one of the following:

1. Recursion can throw up problems with execution time when a function calls itself more than once; these problems can't generally be solved.
2. Recursion can throw up problems with execution time when a function calls itself more than once, but it can be possible to find a clever work-around.
3. Recursion doesn't generally throw up problems with execution time at all.

<!-- grade_id: cell-0ce7f31ea1a61fcb, points: 0, grade: false, solution: true -->
```python
question4f_answer = # (delete, and fill in the value 1, 2, 3 or 4)
### BEGIN SOLUTION
question4f_answer = 2
### END SOLUTION
```

<!-- grade_id: cell-b0e5af91705dd496, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question4f_answer == question4_f())
print('test case passed!')
```

(g) Write and test a recursive implementation of your aliquot chain function, <code>aliquot_chain2</code>, which should terminate at the first occurrence of a number that has occurred before. You're likely to need your <code>factor_sum</code> function again.

Note: this is tricky. One thing to consider is writing, as with part (e), a recursive "inner" function and an "outer" function that simply calls it. In this case, the "inner" function would need to take as its argument not a single value but a list of values; in this way, it has access to information concerning what numbers have already occurred.

<!-- grade_id: cell-717274c920b3892e, points: 0, grade: false, solution: true -->
```python
def factor_sum(a):
    ### BEGIN SOLUTION
    from math import floor, sqrt
    # initalize fsum
    fsum = 0
    # square root of a
    sqrta = floor(sqrt(a))
    # for loop
    for r in range(1,sqrta+1):
        if a%r == 0:
            fsum += r + a//r
    # adjust for perfect squares
    if sqrta ** 2 == a:
        fsum -= sqrta
    return fsum
    ### END SOLUTION

def aliquot_chain_inner(list_so_far):
    ### BEGIN SOLUTION
    # last in list so far
    last_a = list_so_far[-1]
    # base case: number has occurred before
    if last_a in list_so_far[0:-1]:
        return list_so_far
    # recursion step
    else:
        next_a = factor_sum(last_a) - last_a
        return aliquot_chain_inner(list_so_far + [next_a])
    ### END SOLUTION

def aliquot_chain2(a):
    ### BEGIN SOLUTION
    return aliquot_chain_inner([a])
    ### END SOLUTION
```

<!-- grade_id: cell-58522f35705ee5d9, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4g(_globals)
```

## Question 5 (Bonus Question)

The prime 41, can be written as the sum of six consecutive primes:
$$
41 = 2 + 3 + 5 + 7 + 11 + 13
$$
This is the longest sum of consecutive primes that adds to a prime below one hundred. The longest list of consecutive primes below one thousand that adds to a prime contains 21 terms and its sum is 953 (unlike the above example, it doesn't begin with 2).

Write a function which takes an integer `n` and finds the longest chain of consecutive primes whose sum is a prime less than or equal to `n`, and returns (i) the value of this sum and (ii) the length of this chain of primes. For example, if `n = 100`, then the function should return `(41, 6)`.

Sometimes, there will be more than one list of maximum length whose sum is prime; in those circumstances, your code should ideally return the largest of these prime sums (though the autograder will be forgiving if it doesn't!)

**Hint:** Use the sieving method we developed for finding primes. For maximum execution speed (e.g. to try to pass the final, most challenging time test), see if you can develop a numpy version. It is possible, on Phil's now fairly aged laptop, to sieve so quickly that the last test is passed in under half a second!

**Note:** The time limits for the varying sized inputs.

<!-- grade_id: cell-746a820ef58a23b4, points: 0, grade: false, solution: true -->
```python
import numpy as np
def prime_chain(n):    
    """
    For a given n, seeks the longest list of consecutive primes whose sum 
    is a prime less than or equal to n, returning the length of the list
    and the value of the sum
    """
    ### BEGIN SOLUTION
    # calculate primes less than or equal to n
    prime_list = primes(n)
    nprimes = len(prime_list)
    
    # initialize for upper bound calculation
    tot = 0
    counter = 0
    
    # upper bound calculation: find length of longest chain beginning with 2
    # whose sum is less than or equal to n
    while tot <= n:
        tot += prime_list[counter]
        counter += 1
    length_upper_bound = counter - 1
    
    # largest possible total (before testing for primality)
    start_tot = tot - prime_list[counter-1]
    
    # loop backwards over chain lengths, starting with the upper bound
    for chain_length in range(length_upper_bound,0,-1):
        # initialize total
        tot = start_tot
        # loop backwards over all possible starting values
        for start in range(length_upper_bound-chain_length,-1,-1):
            # if it's in our list of primes, return it and the corresponding chain length
            if tot in prime_list:
                return(tot, chain_length)
            # adjust value of tot, dropping the end value and adding the previous
            tot -= prime_list[chain_length+start-1]
            tot += prime_list[start-1]
        # adjust starting total, removing first value
        start_tot -= prime_list[length_upper_bound-chain_length]
    

def primes(n):
    """
    Uses the Sieve of Eratosthenes to return a list of the primes
    less than or equal to n
    """
    
    # initialize: logical array of True
    prime_position_array = np.ones(n+1,dtype=bool)
    
    # values 0 and 1 are False
    prime_position_array[0:2] = False
    
    # iterate up to square root
    for i in range(2,int(np.ceil(np.sqrt(n+1)))+1):
        # if position i is labelled True (the number in this position is prime)...
        if prime_position_array[i]:
            # ... label positions i**2, i**2+i, i**2+2i etc as False ("graying out")
            prime_position_array[i*i:n+1:i] = False
     
    # uncomment the line below to print out the logical array that we end up with
    #print(prime_position_array)
    
    # return numbers corresponding to the value True
    #return np.arange(n+1)[prime_position_array]
    return np.where(prime_position_array)[0]

### END SOLUTION
```

<!-- grade_id: cell-677337a3ba85db4b, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# unit test
# 1 Mark

assert(41 == prime_chain(100)[0])
assert(6 == prime_chain(100)[1])
print('All Tests Passed!!!')

```

<!-- grade_id: cell-809ccc45c2022063, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# small input
# 1 Marks

assert(953 == prime_chain(1000)[0])
assert(21 == prime_chain(1000)[1])
print('All Tests Passed!!!')
```

<!-- grade_id: cell-18ec81e49e89b4ec, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# medium input
#--------------------------
# 15 second time limit
#--------------------------
# 1 Mark
import time
start      = time.time()
time_limit = 15

assert((9521, 65) == prime_chain(10000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')
 
```

<!-- grade_id: cell-d3223e3d75e6b1b6, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# large inputs
#--------------------------
# 30 second time limit
#--------------------------
# 1 Mark

start      = time.time()
time_limit = 30

assert((92951, 183) == prime_chain(100000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```

<!-- grade_id: cell-5257d126ea8742c1, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# large inputs
#--------------------------
# 60 second time limit
#--------------------------
# 1 Mark

start      = time.time()
time_limit = 60

assert((997651, 543) == prime_chain(1000000))
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```

<!-- grade_id: cell-41ef14b1d6c0a03e, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
#--------------------------
# 60 second time limit
#--------------------------
# huge inputs
# 2 Marks

start      = time.time()
time_limit = 60
pc = prime_chain(10000000) 
assert((9951191, 1587) == pc or (9964597, 1587) == pc)
if abs(time.time() - start) < time_limit:
    print('All Tests Passed!!!')
else:
    raise AssertionError('Program Too Slow!!!')

```

```python

```