```python
from autograder_term1week10 import *
_globals = globals()
import matplotlib.pyplot as plt
%matplotlib inline
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Matt Woolway, Boumediene Hamzi</center>

# <center> Worksheet 5: Functions</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-46aa21238875aab2, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

## Question 1

(a) Write and test a function called `xsqdplusone` that takes one argument $x$ (which you can assume to be numerical) and returns the value of $x^2 + 1$. Your function should return a number of the same type as $x$.

<!-- grade_id: cell-b00091df7b05694f, points: 0, grade: false, solution: true -->
```python
def xsqdplusone(x):
    ### BEGIN SOLUTION
    return x**2 + 1
    ### END SOLUTION
```

<!-- grade_id: cell-35116d17e85ec49c, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(xsqdplusone(3+2j) == (6+12j))
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-9f3ccf3b0dcc3a1a, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

(b) Write and test a function called <code>hailstone_step</code> that takes one argument $a$ (which you can assume to be a positive integer) and returns an integer equal to $a/2$ if $a$ is even, and $3a+1$ if $a$ is odd.

<!-- grade_id: cell-7582dbb0352598f5, points: 0, grade: false, solution: true -->
```python
def hailstone_step(a):
    ### BEGIN SOLUTION
    if a % 2 == 0:
        return a//2
    else:
        return 3*a+1
    ### END SOLUTION
```

<!-- grade_id: cell-5d58e3855945bf93, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(hailstone_step(7) == 22)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-03a43bd44d34aae1, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

(c) Write a function called <code>my_factorial</code> that takes one argument $n$ (which you can assume to be a non-negative integer) and returns its factorial, $n!$.

<!-- grade_id: cell-72ce8a0b51ff800f, points: 0, grade: false, solution: true -->
```python
def my_factorial(n):
    ### BEGIN SOLUTION
    # initalize return value facn
    facn = 1
    # for loop, for r = 2, 3, ... n
    for r in range(2,n+1):
        # multiply facn by r
        facn *= r
    return facn
    ### END SOLUTION
```

<!-- grade_id: cell-3bb5ea38caa35539, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(my_factorial(5) == 120)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases, including edge cases.')
```

<!-- grade_id: cell-12fd18ed347a5d4e, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Write and test a function called <code>fibonacci</code>, which takes as its argument $n$, which you can assume to be a positive integer, and returns the $n$th Fibonacci number. (Note: the first Fibonacci number can be taken to be 1, the second 1, the third 2, etc, so
```python
fibonacci(5)
```
should return 5.)

<!-- grade_id: cell-083814268f826c5d, points: 0, grade: false, solution: true -->
```python
def fibonacci(n):
    ### BEGIN SOLUTION
    # initalize a and b
    a, b = 1, 1
    # for loop
    for r in range(3,n+1):
        # update a and b using Fibonacci relation
        a, b = b, a+b
    return b
    ### END SOLUTION
```

<!-- grade_id: cell-0e433bf224e67d1a, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(fibonacci(5) == 5)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-512e72d91dabe64f, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

(e) Write and test a function called <code>factor_sum</code> that takes as its argument <code>a</code>, which you can assume to be a positive integer, and returns the sum of all its factors including itself. Try to make your function reasonably efficient, without too many unnecessary iterations.

<!-- grade_id: cell-ed953a434d1a00e3, points: 0, grade: false, solution: true -->
```python
def factor_sum(a):
    ### BEGIN SOLUTION
    from math import floor, sqrt
    # initalize fsum
    fsum = 0
    # square root of a
    sqrta = floor(sqrt(a))
    # for loop
    for r in range(1,sqrta+1):
        if a%r == 0:
            fsum += r + a//r
    # adjust for perfect squares
    if sqrta ** 2 == a:
        fsum -= sqrta
    return fsum
    ### END SOLUTION
```

<!-- grade_id: cell-9124bb3900fcaf38, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
assert(factor_sum(24) == 60)
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-b19e110c9aa2d269, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1e(_globals)
```

(f) Write a function called <code>nest</code>, which takes as its arguments a function <code>f</code>, a starting value <code>x0</code> and a number of iterations <code>n</code>, and returns 

<center>
    <code>f(f(f(...f(x0)...))),</code>
</center>

the <code>n</code>th iterate of <code>f</code> starting at <code>x0</code>.

Test your function. The following code should return the value 

$$\cos(\cos(\cos(\cos(\cos(1.0))))) = 0.7013687736227566,$$ 

for example (though you should test more than once).

```python
from math import cos
nest(cos, 1.0, 5)
```

Using your function, calculate the 20th iterate of $f$ starting at $0.3$, where

$$f(x) = 4\,x\,(1-x).$$

Make sure to pass this to the variable <code>x20</code>.

(Hint; for the last part, you may wish to use something called a <b>lambda-expression</b>; take a sneak peek at the notes for Term 2, Week 1. A lambda-expression is not required for this question, but you might find it convenient.)

<!-- grade_id: cell-f036c883340009fc, points: 0, grade: false, solution: true -->
```python
def nest(f, x0, n):
    ### BEGIN SOLUTION
    # initalize x
    x = x0
    # for loop
    for r in range(n):
        # update x
        x = f(x)
    return x

x20 = nest(lambda x: 4*x*(1-x), 0.3, 20)
    ### END SOLUTION

x20 = # delete, and fill in your code
print(x20)
```

<!-- grade_id: cell-eb91dac56e3e4ea7, points: 0, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Unit test
# 0 Marks
from math import cos
from numpy import isclose
assert(isclose(nest(cos, 1.0, 5), 0.7013687736227566))
print('Unit test passed! Run some of your own tests, making sure to try a variety of cases.')
```

<!-- grade_id: cell-8dda4608d9bf0aec, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1f(_globals)
```

## Question 2


(a) Write and test a function called <code>hailstone_sequence</code> which takes as its arguments a starting value <code>a</code> and a number of iterations <code>n</code>, both assumed to be positive integers. It should then return, as a list, the sequence generated by iterating "if <code>a</code> is even, halve it; otherwise, multiply it by 3 and add 1" <code>n</code> times; the list should be of length <code>n+1</code>.

(Hint: you may wish to call the function you've already written called <code>hailstone_step</code>.)

<!-- grade_id: cell-bf69a919a8656af0, points: 0, grade: false, solution: true -->
```python
def hailstone_sequence(a, n):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    # for loop
    for r in range(n):
        # iterate and append
        a = hailstone_step(a)
        alist.append(a)
    return alist
    ### END SOLUTION
```

<!-- grade_id: cell-bad32cb1372c0531, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

(b) Write and test a function called <code>fibonacci_sequence</code> which takes as its argument $n$ (which you can assume to be a positive integer) and returns a list consisting of the Fibonacci sequence up to the $n$th term. The first Fibonacci number should be treated as 1, and the second as 1, so that
```python
fibonacci_sequence(5)
```
should return
```
[1, 1, 2, 3, 5]
```

<!-- grade_id: cell-819fd6713dfb2bbd, points: 0, grade: false, solution: true -->
```python
def fibonacci_sequence(n): 
    ### BEGIN SOLUTION
    # case n=1
    if n==1:
        return [1]
    else:
        # initialize a, b and fiblist
        a, b = 1, 1
        fiblist = [a, b]
        # for loop
        for r in range(3, n+1):
            # update and append
            a, b = b, a + b
            fiblist.append(b)
        return fiblist
    ### END SOLUTION
```

<!-- grade_id: cell-3e5dee79d8ab2138, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2b(_globals)
```

(c) Write a function called <code>aliquot_chain</code> which takes as its argument a starting value <code>a</code> and a number of iterations <code>n</code>, both non-negative integers. It should return a list of length <code>n+1</code>, starting with <code>a</code>, where each number in the list (starting from the second) is the sum of the factors of the previous number <em>not including that number itself</em>.

Test your function: the code

<code>
aliquot_chain(24, 5)
</code>

should return the value <code>[24, 36, 55, 17, 1, 0]</code>, for example.

<!-- grade_id: cell-39b4848ed884388d, points: 0, grade: false, solution: true -->
```python
def aliquot_chain(a, n):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    for r in range(n):
        # update and append
        a = factor_sum(a) - a
        alist.append(a)
    return alist
    ### END SOLUTION
```

<!-- grade_id: cell-4eaaee03a4b30bee, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

(d) Write and test a function called <code>nest_list</code> which takes as its arguments a function <code>f</code>, a starting value <code>x0</code> and a number of iterations <code>n</code>, and returns the list

<center>
    <code>[x0, f(x0), f(f(x0)), ..., f(f(f(...f(x0)...)))],</code>
</center>

(that is, all <code>n+1</code> iterates of <code>f</code>, from the starting value <code>x0</code> to the <code>n</code>th).

Use your function to calculate the terms in the sequence generated by 20 iterations of $f$, where

$$f(x) = 4\,x\,(1-x),$$

starting from $x=0.3$. Pass this to the variable <code>xlist20</code>.

<!-- grade_id: cell-2ae8066aa476fbc7, points: 0, grade: false, solution: true -->
```python
def nest_list(f, x0, n):
    ### BEGIN SOLUTION
    # initialize x and xlist
    x = x0
    xlist = [x]
    # for loop
    for r in range(n):
        # update and append
        x = f(x)
        xlist.append(x)
    return xlist

xlist20 = nest_list(lambda x: 4*x*(1-x), 0.3, 20)
    ### END SOLUTION
```

<!-- grade_id: cell-6ecd8661ad6362cf, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2d(_globals)
```

(e) Write a function called `duffing_iterates` which takes as its arguments two parameters `a` and `b`, two starting values `x0` and `y0`, and a number of iterations `n`, and returns a 2-tuple containing the lists `[x0, x1, x2, ..., xn]` and `[y0, y1, y2, ..., yn]` generated by the Duffing map
$$x_{n+1}=y_n,$$
$$y_{n+1}=-b\,x_n+a\,y_n-{y_n}^3.$$

Check your function; check in particular that it is capable of assigning values to two variables simultaneously.

Use the `plot` function to generate a 2D point plot of the iterates in the case $a=2.75$, $b=0.2$, $x_0=y_0=0.5$ and $n=100000$.

<!-- grade_id: cell-5fa7ae48378128a4, points: 0, grade: false, solution: true -->
```python
def duffing_iterates(a, b, x0, y0, n):
    ### BEGIN SOLUTION
    # initialize x, y, xlist, ylist
    x, y = x0, y0
    xlist, ylist = [x], [y]
    # for loop
    for r in range(n):
        # update and append
        x, y = y, -b*x + a*y - y**3
        xlist.append(x)
        ylist.append(y)
    return xlist, ylist

its = duffing_iterates(2.75, 0.2, 0.5, 0.5, 100000)
plt.plot(its[0], its[1], '.', markersize=0.1)
    ### END SOLUTION
```

<!-- grade_id: cell-5ce51157eb9854dc, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2e(_globals)
```

## Question 3

(a) Write and test a version of your hailstone sequence function, called <code>hailstone_sequence2</code>, which takes only the argument <code>a</code> and which terminates not after a set number of iterations but when the sequence hits the value 1. (You probably want to use a <code>while</code> loop.)

<!-- grade_id: cell-48e58cfbd11e3e3b, points: 0, grade: false, solution: true -->
```python
def hailstone_sequence2(a):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    # while loop
    while a > 1:
        # iterate and append
        a = hailstone_step(a)
        alist.append(a)
    return alist
    ### END SOLUTION
```

<!-- grade_id: cell-1c2b062449dc979c, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

(b) Write a version of your aliquot chain function, called <code>aliquot_chain2</code>, which takes only the argument <code>a</code> and which terminates not after a set number of iterations but when the sequence encounters a value it has encountered before. 

Test your function: <code>aliquot_chain2(24)</code> should return <code>[24, 36, 55, 17, 1, 0, 0]</code>, whereas <code>aliquot_chain2(28)</code> should return <code>[28, 28]</code>, and <code>aliquot_chain2(220)</code> should return <code>[220, 284, 220]</code>.

<!-- grade_id: cell-1b189024c6e59d88, points: 0, grade: false, solution: true -->
```python
def aliquot_chain2(a):
    ### BEGIN SOLUTION
    # initialize alist
    alist = [a]
    while not a in alist[0:-1]:
        # update and append
        a = factor_sum(a) - a
        alist.append(a)
    return alist
    ### END SOLUTION
```

<!-- grade_id: cell-425a664ab09929e4, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

(c) Write and test a function called <code>fixed_point_list</code>, which takes as its input a function <code>f</code>, an initial value <code>x0</code> and a <code>tolerance</code>, and returns a list consisting of 

<code>[x0, f(x0), f(f(x0)), ...]</code>, 

which terminates when successive iterates are within <code>tolerance</code> of one another. (Make sure you test it only on iterations that you know converge; later, we'll look at how to build in safeguards to prevent infinite loops when an iteration fails to converge.)

<!-- grade_id: cell-70aad741029c978d, points: 0, grade: false, solution: true -->
```python
def fixed_point_list(f, x0, tolerance):
    ### BEGIN SOLUTION
    # initialize oldx, newx and xlist
    oldx = x0
    newx = f(oldx)
    xlist = [oldx, newx]
    # while loop
    while abs(oldx-newx) > tolerance:
        # update and append
        oldx = newx
        newx = f(oldx)
        xlist.append(newx)
    return xlist
    ### END SOLUTION
```

<!-- grade_id: cell-753024b1ab5e656d, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3c(_globals)
```

(d) Write and test a function called <code>binary_digits</code>, which takes as its input a non-negative int $a$, and returns its binary digits as a list of ones and zeros. The list should begin with the first non-zero digit, unless $a=0$ in which case the list should contain a single zero.

<!-- grade_id: cell-a11a1e7e8036b571, points: 0, grade: false, solution: true -->
```python
def binary_digits(a):
    ### BEGIN SOLUTION
    # edge case
    if a==0:
        return [0]
    # initialize list
    digits = []
    # while loop
    while a>0:
        digits.append(a%2)
        a //= 2
    digits.reverse()
    return digits
    ### END SOLUTION
```

<!-- grade_id: cell-c6c0d9e08319602c, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3d(_globals)
```

Output:
```
Test 1: a is a power of 2

binary_digits(2**15) returned the correct list of digits
The digits are in the correct order

Test 2: a is not a power of 2

binary_digits(14827) returned the correct list of digits
The digits are in the correct order

Test 3: edge cases

binary_digits(1) returned the correct list of digits
binary_digits(0) returned the correct list of digits

3 marks out of 3
3
```

(e) Write and test a function called <code>binary_digits_float</code>, which takes as its input a non-negative float $a$, and a non-negative int $n$, and returns its binary digits as two lists of ones and zeros. The first list should consist of the binary expansion of the integer part of $a$, and the second should consist of the binary expansion of its fractional part to $n$ (binary) places (just truncate this list; don't bother with rounding). 

<!-- grade_id: cell-9dc8503d972edd83, points: 0, grade: false, solution: true -->
```python
def binary_digits_float(a, n):
    ### BEGIN SOLUTION
    # integer and fractional parts
    intpart = int(a)
    fracpart = a - intpart
    
    # binary expansion of integer part
    list0 = binary_digits(intpart)
    # binary expansion of fractional part to n places, truncated
    list1 = binary_digits(int(2**n*fracpart))
    
    # possible bug is that expansion begins with zeros that aren't returned
    # by binary_digits, so we pad with zeros at the beginning
    list1 = [0 for r in range(n-len(list1))] + list1
    
    return list0, list1    
    ### END SOLUTION
```

<!-- grade_id: cell-51cbb375d8c2f076, points: 4, grade: true, solution: false -->
```python
# 4 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3e(_globals)
```

Output:
```
Test 1: typical case

binary_digits_float(12.8191,10) returned the correct list of digits for the integer part
binary_digits_float(12.8191,10) returned the correct list of digits for the fractional part

Tests 2: fractional part begins with zeros

binary_digits_float(12.2191,10) returned the correct list of digits for the fractional part,
including the two leading zeros

Test 3: integer part is zero

binary_digits_float(0.8191,10) returned the correct list of digits for the integer part
binary_digits_float(0.8191,10) returned the correct list of digits for the fractional part

Test 4: fractional part is zero

binary_digits_float(12.0,10) returned the correct list of digits for the integer part
binary_digits_float(12.0,10) returned the correct list of digits for the fractional part

4 marks out of 4
4
```

## Question 4

The method of calculating the primes by successive filtering is called the <b>sieve of Eratosthenes</b>; Earlier, you explored an implementation of this based on <em>conditional comprehensions</em>. You started with the integers between 2 and 400 inclusive; you then removed all the multiples of 2 except 2 itself; then all the multiples of 3 except 3 itself; then all the multiples of 5 except 5 itself; and so on, using primes each time.

Write and test a function called <code>my_primes</code> that takes as its argument a single <code>n</code>, which you can assume to be an integer greater than 1, and, using the sieve of Eratosthenes, returns a list of all the primes between 2 and <code>n</code> inclusive. Make your function reasonably efficient; don't put in lots of unnecessary iterations.

<!-- grade_id: cell-09d7c2ad10004e0e, points: 0, grade: false, solution: true -->
```python
def my_primes(n):
    ### BEGIN SOLUTION
    # initialize ints, r, p
    ints = list(range(2,n+1))
    r = 0
    p = ints[r]
    # while loop
    while p**2 <= ints[-1]:
        # filter
        ints = [a for a in ints if a == p or a % p > 0]
        # increment index
        r += 1
        # next prime
        p = ints[r]
    return ints
    ### END SOLUTION
```

<!-- grade_id: cell-c749e7603ceb65e9, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4(_globals)
```

## Bonus Questions

### Question 1

Write a function, `symCheck`, that takes as input a matrix, expressed as a list of lists, and determines whether or not the matrix is symmetric. The function should return the strings `not symmetric` or `symmetric`. You may not use any obvious builtin function.

<!-- grade_id: cell-f3f2aefc41f2911d, points: 0, grade: false, solution: true -->
```python
# Question 1
def symCheck(A):
    ### BEGIN SOLUTION
    import numpy as np
    A    = np.asarray(A)
    r, c = A.shape
    if r != c:
        return 'not symmetric'
    for rows in range(r):
        if any(A[rows, :] != A[:, rows]):
            return 'not symmetric'
    else:
        return 'symmetric'
    ### END SOLUTION
```

<!-- grade_id: cell-0ac2746c8242a091, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3]]
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-af2c44662743c74d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3], [2, 1, 3]]
assert(symCheck(A) == 'not symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-35974a0e424c744f, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
A = [[1, 2, 3], [2, 1, 3], [3, 3, 3]]
assert(symCheck(A) == 'symmetric')
print('Test case passed!!!')
```

<!-- grade_id: cell-3c74338d7cde78dc, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
### BEGIN HIDDEN TESTS
import numpy as np
b = np.random.randint(-100, 100, size = (10, 10))
test = (b + b.T)
assert(symCheck(test) == 'symmetric')
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

<!-- grade_id: cell-6a3ea90ee9b77c8a, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
### BEGIN HIDDEN TESTS
import numpy as np
b = np.random.randint(-100, 100, size = (10, 10))
assert(symCheck(b) == 'not symmetric')
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

### Question 2

Write a function that takes in two arrays and finds the number of integers that satisfy the following conditions:

* All elements of the first array are factors of the integer under consideration.
* The integer under consideration is a factor of all the elements in the second array.
    
For example, if $a = [2, 6]$ and $b = [24, 36]$, there are 2 numbers between them which satisfy the above, namely 6 and 12. To check, consider the first value 6: 6 % 2 = 0, 6 % 6 = 0, as well as 24 % 6 = 0 and 36 % 6 = 0. Similarly, 12 % 2 = 0, 12 % 6 = 0,  24 % 12 = 0 and 36 % 12 = 0.

You may use the `math` function `gcd`, which calculates greatest common divisors.

<!-- grade_id: cell-c1ca2727aaac882f, points: 0, grade: false, solution: true -->
```python
def setDiff(a, b):
    ### BEGIN SOLUTION
    from functools import reduce
    from math import gcd
    counter = 0
    lcmNum = reduce(lcm, a)
    gcdNum = reduce(gcd, b)
    mults  = lcmNum
    while mults <= gcdNum:
        print(gcdNum, mults, gcdNum % mults)
        if gcdNum % mults == 0:
            counter += 1
        mults += lcmNum
    return counter

def lcm(a, b):
    from math import gcd
    return a * b // gcd(a, b)
    ### END SOLUTION
```

```python
setDiff([4,6],[360,480])
```

<!-- grade_id: cell-aca1bbc826cd86df, points: 2, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 2 marks
# Unit test
bonus_q2(_globals)
```

<!-- grade_id: cell-bb1274933762d2a4, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
### BEGIN HIDDEN TESTS
a = [2, 4]
b = [16, 32, 96]
assert(setDiff(a, b) == 3)
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

<!-- grade_id: cell-9f27c0f689941e99, points: 1, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 1 mark
### BEGIN HIDDEN TESTS
a = [1]
b = [72, 48]
assert(setDiff(a, b) == 8)
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

```python

```