```python
from autograder_term2week4 import *
_globals = globals()
import numpy as np
import random as rnd
import numpy.random as nrnd
import matplotlib.pyplot as plt
from mpl_toolkits.mplot3d import Axes3D
%matplotlib inline
```

Output:
```
Autograder loaded successfully!
Remember to always restart and run all from the Kernel menu before submitting!
```

# <center> Introduction to Mathematical Computing </center>
### <center>Phil Ramsden, Boumediene Hamzi, Sam Brzezicki</center>

# <center> Worksheet 8: random, numpy.random and matplotlib.pyplot</center>

**NOTE:** When answering a question, remember to remove the `raise NotImplementedError()` line.

## Question 0: don't miss this bit!

In the file called `my_username.txt`, type your <b>short-form College user name</b> (this will usually be one or more letters followed by three or more figures, such as qk4019). Don't use quote marks. Then save the file.

<!-- grade_id: cell-e651f6a984190fff, points: 5, grade: true, solution: false -->
```python
# 5 Marks
# Do not try to delete this cell
# Run this cell for grading of Question 0
_globals = globals()
question0(_globals)
```

Output:
```
You've changed the contents of the file; thank you!
Your username has been recognised; thank you!

5 out of 5 marks
5
```

## Question 1

(a) Using the `random` function from the `random` module in a comprehension, or otherwise, create a list of 10 random floats between 0 and 1, calling it `randlist10`.

<!-- grade_id: cell-8effae7d5d727daf, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
randlist10 = [rnd.random() for n in range(10)]
### END SOLUTION
print(randlist10)
```

<!-- grade_id: cell-3bb8a366d0cd2f16, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1a(_globals)
```

(b) Using a single call to the `random` function from the `numpy.random` submodule (a comprehension or loop shouldn't be necessary), create an <em>array</em> of ten random floats between 0 and 1, calling it `randarray10`.

<!-- grade_id: cell-8aa73d4f6c33fc13, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
randarray10 = nrnd.random(10)
### END SOLUTION
print(randarray10)
```

<!-- grade_id: cell-ce170d2b98f8e75d, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1b(_globals)
```

(c) Using a single call to the `random` function from the `numpy.random` submodule, create a 20 by 10 array of random floats between 0 and 1, calling it `randarray20by10`.

<!-- grade_id: cell-5af25bc81361e627, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
randarray20by10 = nrnd.random([20,10])
### END SOLUTION
print(randarray20by10)
```

<!-- grade_id: cell-fca540c130798a46, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1c(_globals)
```

(d) Using the `mean` function from `numpy`, with the keyword argument `axis` set to 1, calculate the 20 means of the 20 rows of your 20 by 10 array, as an array called `mean20`.

<!-- grade_id: cell-afba63ef39f0040f, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
mean20 = np.mean(randarray20by10,axis=1)
### END SOLUTION
print(mean20)
```

<!-- grade_id: cell-2a50ad10a36fbb2e, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1d(_globals)
```

(e) Write and test a function called `rowmeans` which takes as its arguments two positive integers `nrows` and `ncols`, and returns an array with `nrows` elements, consisting of the means of each of the rows of an `nrows` by `ncols` array of random floats between 0 and 1.

<!-- grade_id: cell-fd01131afb84211e, points: 0, grade: false, solution: true -->
```python
def rowmeans(nrows, ncols):
    ### BEGIN SOLUTION
    from numpy import mean
    from numpy.random import random
    return mean(random([nrows,ncols]),axis=1)
    ### END SOLUTION
```

<!-- grade_id: cell-38309cec0bda6466, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question1e(_globals)
```

(f) Write and test a function called `rowmeanshist` which takes as its arguments two positive integers `nrows` and `ncols`, and generates a histogram of `rowmeans(nrows, ncols)`.

<!-- grade_id: cell-bd17254127375e4a, points: 0, grade: false, solution: true -->
```python
def rowmeanhist(nrows, ncols):
    ### BEGIN SOLUTION
    plt.hist(rowmeans(nrows,ncols))
    ### END SOLUTION
```

(g) Using your function investigate the distributions of the means for `nrows` equal to 400 and `ncols` equal to 1, 2, 4, 8 and 16.

Which of the following statements is true?

I. As the number of columns increases, the mean tends to increase.<br/>
II. As the number of columns increases, the means tend to vary less.<br/>
III. As the number of columns increases, the distribution of the means becomes more approximately normal.<br/>

1. All of them.
2. All of them except I.
3. All of them except II.
4. All of them except III.

<!-- grade_id: cell-48bf614b9017ac21, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
question1g_answer = 2
### END SOLUTION
question1g_answer = # insert one of 1, 2, 3 or 4

```

<!-- grade_id: cell-bb217909883073ea, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question1g_answer == question1g())
print('test case passed!')
```

## Question 2

(a) Define both `x_values1` and `y_values1` as 1D NumPy arrays of 25 equally-spaced floats between $-1.2$ and $1.2$ inclusive, and use the `meshgrid` function to define `x1` and `y1` as 25-by-25 arrays of, respectively, $x$- and $y$-coordinates specifying a lattice in the plane.

<!-- grade_id: cell-08956e751f87f27e, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
x_values1 = np.linspace(-1.2, 1.2, 25)
y_values1 = np.linspace(-1.2, 1.2, 25)
x1, y1 = np.meshgrid(x_values1, y_values1)
### END SOLUTION

# I don't think it's worth printing these values; you may disagree!
```

<!-- grade_id: cell-6f64349f1b251f07, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2a(_globals)
```

(b) Use the `imshow` function from `matplotlib.pyplot` to create, on two separate pairs of axes, visualisations of the functions

$$u(x,y) = x^4-6\,x^2\,y^2+y^4$$
and
$$v(x,y) = 4\,x^3\,y-4\,x\,y^3$$
on the points of this lattice.

(Autograding won't be able to help much here.)

[Mathematical note: these are, respectively, the real and imaginary parts of $(x+i\,y)^4$, where $x$ and $y$ are real.]

<!-- grade_id: cell-999f26a0d59f1ac3, points: 0, grade: false, solution: true -->
```python
# Cell for first plot
### BEGIN SOLUTION
u1 = x1**4 - 6*x1**2*y1**2+y1**4
plt.imshow(u1)
### END SOLUTION
```

<!-- grade_id: cell-b9e9238ad749ec25, points: 0, grade: false, solution: true -->
```python
# Cell for second plot
### BEGIN SOLUTION
v1 = 4*x1**3*y1 - 4*x1*y1**3
plt.imshow(v1)
### END SOLUTION
```

(c) Repeat parts (a) and (b), but this time use `x_values2` and `y_values2`, which you should define as 1D arrays of 121 equally-spaced floats between $-1.2$ and $1.2$ inclusive (and then go on to define `x2` and `y2` accordingly). This should create higher-resolution plots. 

(Again, autograding won't help you with the plots!)

<!-- grade_id: cell-3f2ed706aefec34b, points: 0, grade: false, solution: true -->
```python
# Cell for setting up variables
### BEGIN SOLUTION
x_values2 = np.linspace(-1.2, 1.2, 121)
y_values2 = np.linspace(-1.2, 1.2, 121)
x2, y2 = np.meshgrid(x_values2, y_values2)
### END SOLUTION

# I don't think it's worth printing these values; you may disagree!
```

<!-- grade_id: cell-650df88834e1db45, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question2c(_globals)
```

<!-- grade_id: cell-f93ed31b5fdc38b6, points: 0, grade: false, solution: true -->
```python
# Cell for first plot
### BEGIN SOLUTION
u2 = x2**4 - 6*x2**2*y2**2+y2**4
plt.imshow(u2)
### END SOLUTION
```

<!-- grade_id: cell-58cb28ff7d8687c1, points: 0, grade: false, solution: true -->
```python
# Cell for second plot
### BEGIN SOLUTION
v2 = 4*x2**3*y2 - 4*x2*y2**3
plt.imshow(v2)
### END SOLUTION
```

(d) Create, on the same pair of axes, contour plots of $u(x,y)$ (in blue) and $v(x,y)$ (in red). Use `x2` and `y2`, and make the scales on both axes equal.

<!-- grade_id: cell-b63bc1e74636097d, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
plt.contour(x2,y2,u2,colors='blue')
plt.contour(x2,y2,v2,colors='red')
plt.axis('equal');
### END SOLUTION
```

(e) Which of the following is consistent with your diagram?

<ol>
    <li>The contours of ${\rm Re}\left((x+i\,y)^4\right)$ and ${\rm Im}\left((x+i\,y)^4\right)$ intersect at right angles everywhere.</li>
    <li>The contours of ${\rm Re}\left((x+i\,y)^4\right)$ and ${\rm Im}\left((x+i\,y)^4\right)$ intersect at right angles everywhere except at $(x+i\,y)=0$.</li>
    <li>The contours of ${\rm Re}\left((x+i\,y)^4\right)$ and ${\rm Im}\left((x+i\,y)^4\right)$ do not in general intersect at right angles.</li>
</ol>

<!-- grade_id: cell-336d594baf119478, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
question2e_answer = 2
### END SOLUTION
question2e_answer = # insert one of 1, 2 or 3
```

<!-- grade_id: cell-aadae8f086ef3bde, points: 1, grade: true, solution: false -->
```python
# 1 Mark
assert(question2e_answer == question2e())
print('test case passed!')
```

(f) Create, on the same pair of axes, a diagram consisting of the contours of $v(x,y)$, obtained using `x2` and `y2`, on which is superimposed a quiver plot of the vector field 
$$\left(\frac{\partial v}{\partial x},\frac{\partial v}{\partial y}\right),$$
obtained using the coarser-grained `x1` and `y1`. Show the former in blue and the latter in red. Make the scale the same on both axes.

[Note: it's probably easiest to do the differentiations by hand, though you're welcome to use SymPy if you like.]

<!-- grade_id: cell-b0554a745e8b4bdb, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
dvdx1 = 12*x1**2*y1-4*y1**3
dvdy1 = 4*x1**3-12*x1*y1**2
plt.contour(x2,y2,v2,colors='blue')
plt.quiver(x1,y1,dvdx1,dvdy1,color='red')
plt.axis('equal');
### END SOLUTION
```

(g) Create, on the same pair of axes, a diagram consisting of the contours of $v(x,y)$, obtained using `x2` and `y2`, on which is superimposed a quiver plot of the vector field 
$$\left(\frac{\partial u}{\partial x},\frac{\partial u}{\partial y}\right),$$
obtained using the coarser-grained `x1` and `y1`. Show the former in blue and the latter in red. Make the scale the same on both axes.

[Note: this isn't a typo; I really do mean to use $v$ for the contour plot and $u$ for the quiver plot here.]

<!-- grade_id: cell-c3cc1ae42fcc34a9, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
dudx1 = 4*x1**3-12*x1*y1**2
dudy1 = -12*x1**2*y1+4*y1**3
plt.contour(x2,y2,v2,colors='blue')
plt.quiver(x1,y1,dudx1,dudy1,color='red')
plt.axis('equal');
### END SOLUTION
```

(h) Create on separate sets of axes, using 3D projections, a wireframe plot, a surface plot and a 3D contour plot of $v(x, y)$.  In each case, use the finer-grained versions of the data.

<!-- grade_id: cell-61d50292422b38ae, points: 0, grade: false, solution: true -->
```python
# Cell for first plot
### BEGIN SOLUTION
fig = plt.figure()
ax = plt.axes(projection='3d')
ax.plot_wireframe(x2, y2, v2)
### END SOLUTION
```

<!-- grade_id: cell-4610a5d614edfa07, points: 0, grade: false, solution: true -->
```python
# Cell for second plot
### BEGIN SOLUTION
fig = plt.figure()
ax = plt.axes(projection='3d')
ax.plot_surface(x2, y2, v2)
### END SOLUTION
```

<!-- grade_id: cell-91bd12f9150c8117, points: 0, grade: false, solution: true -->
```python
# Cell for third plot
### BEGIN SOLUTION
fig = plt.figure()
ax = plt.axes(projection='3d')
ax.contour(x2, y2, v2)
### END SOLUTION
```

## Question 3

(a) Define
```python
n = 10
p = 0.3
```

Then define `r` as a 1D array consisting of the integers between 0 and `n` inclusive.

<!-- grade_id: cell-0d14e655b59b9cee, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
n = 10
p = 0.3
r = np.arange(0,n+1)
### END SOLUTION
print(n)
print(p)
print(r)
```

<!-- grade_id: cell-de68806389692a9b, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3a(_globals)
```

(b) Using `numpy`'s `cumprod` function, or otherwise, define `rfac` as a 1D array consisting of the integers $0!, 1!, 2!, \dots n!$, and `nmrfac` as a 1D array consisting of the integers $n!, (n-1)!, (n-2)!, \dots, 0!$.

<!-- grade_id: cell-5b04f8de66d3d075, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
rfac = np.append([1],np.cumprod(r[1:]))
nmrfac = rfac[-1::-1]
### END SOLUTION
print(rfac)
print(nmrfac)
```

<!-- grade_id: cell-ccee552c5ba065e3, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3b(_globals)
```

(c) Using `rfac` and `nmrfac`, or otherwise, define `bincoef` as a 1D array consisting of the binomial coefficients

$$\frac{n!}{0!\,n!},\frac{n!}{1!\,(n-1)!},\frac{n!}{2!\,(n-2)!},\dots,\frac{n!}{n!\,0!}.$$

Then define `binprob` as a 1D array of floats, consisting of all the values of
$$\frac{n!}{r!\,(n-r)!}\,p^r\,(1-p)^{n-r}$$
for $r = 0,1,2,\dots,n.$

<!-- grade_id: cell-150814d418bbdee1, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
bincoef = rfac[-1]/(rfac*nmrfac)
binprob = bincoef * p**r * (1-p)**(n-r)
### END SOLUTION
print(bincoef)
print(binprob)
```

<!-- grade_id: cell-079d267bcc0da88a, points: 2, grade: true, solution: false -->
```python
# 2 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3c(_globals)
```

(d) Create, on the same pair of axes, a blue line plot of `binprob` (on the vertical axis) against `r` (on the horizontal axis), a blue point plot of `prob` against `r` (with a nice big marker size like 10), and a red bar chart of `prob` against `r`. 

<!-- grade_id: cell-0e8447c361379dea, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
plt.plot(r,binprob,'b')
plt.plot(r,binprob,'.b',markersize=10)
plt.bar(r,binprob,color='red');
### END SOLUTION
```

(e) Use the `binomial` function from NumPy's `random` submodule to create a sample `data100` of size 100 from the binomial distribution $B(n,p)$, with the above values of $n$ and $p$. Create, on the same pair of axes, a blue line plot of `binprob` (on the vertical axis) against `r` (on the horizontal axis), a blue point plot of `prob` against `r` (with a nice big marker size like 10), and a red <b>density histogram</b> of this data, using bin boundaries $-0.5, 0.5, 1.5, \dots ,10.5$.

Repeat for sample sizes of 1000 and 10000, calling your data `data1000` and `data10000`.

<!-- grade_id: cell-e5d83668406c6e7f, points: 0, grade: false, solution: true -->
```python
# Cell for first plot
### BEGIN SOLUTION
data100 = nrnd.binomial(n,p,100)
plt.plot(r,binprob,'b')
plt.plot(r,binprob,'.b',markersize=10)
plt.hist(data100, bins = np.arange(-0.5,11.0), color='red', density=True);
### END SOLUTION
```

<!-- grade_id: cell-77f317c15e983e2a, points: 1, grade: true, solution: false -->
```python
# For checking data100
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3e1(_globals)
```

<!-- grade_id: cell-81f94ed47c843470, points: 0, grade: false, solution: true -->
```python
# Cell for second plot
### BEGIN SOLUTION
data1000 = nrnd.binomial(n,p,1000)
plt.plot(r,binprob,'b')
plt.plot(r,binprob,'.b',markersize=10)
plt.hist(data1000, bins = np.arange(-0.5,11.0), color='red', density=True);
### END SOLUTION
```

<!-- grade_id: cell-33fb77e5b1e40209, points: 1, grade: true, solution: false -->
```python
# For checking data1000
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3e2(_globals)
```

<!-- grade_id: cell-1525f678157c898f, points: 0, grade: false, solution: true -->
```python
# Cell for third plot
### BEGIN SOLUTION
data10000 = nrnd.binomial(n,p,10000)
plt.plot(r,binprob,'b')
plt.plot(r,binprob,'.b',markersize=10)
plt.hist(data10000, bins = np.arange(-0.5,11.0), color='red', density=True);
### END SOLUTION
```

<!-- grade_id: cell-a5fdc908c5c754fb, points: 1, grade: true, solution: false -->
```python
# For checking data10000
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question3e3(_globals)
```

# Question 4

This question is about <b>iterated function systems</b>. An iterated function system consists of a finite list of functions from the plane to itself (or, more generally, from $n$-dimensional space to itself). A function is chosen at random from the list, and a starting pair of coordinates fed into it, produding a new pair of coordinates. This is fed into a freshly-chosen random function from the list, and so on, for a large number of iterations.

In all the cases we shall consider, the functions will be <b>affine</b>: that is, each will consist of a linear transformation followed by a translation. The system can therefore be completely specified by $n$ transformation matrices and $n$ translation vectors.

(a) Write and test a function called `ifs_step`. Your function should take as its arguments:
<ul>
    <li>a 3D NumPy array called <code>matrices</code>; the first dimension of this array will be some positive integer (let's call it $n$), and the other two dimensions will be 2; this array can therefore represent a list of $n$ matrices;</li>
    <li>a 2D NumPy array called <code>vectors</code>; the first dimension will be the same as that of <code>matrices</code>, and the second will be 2; this array can therefore represent a list of $n$ vectors;</li>
    <li>a 1D NumPy array of length 2 called <code>coords</code>.
</ul>

It should then:
<ul>
    <li>choose a random integer, <code>index</code>, between $0$ and $(n-1)$ inclusive;</li>
    <li>set <code>mat</code> equal to the $(2\times2)$ array given by <code>matrices[index,:,:]</code>;</li>
    <li>set <code>vec</code> equal to the 1D array given by <code>vectors[index,:]</code>;</li>
    <li>return the value of <code>dot(mat, coords) + vec</code>.</li>
</ul>

To test your function, try calculating
```python
ifs_step(mats, vecs, np.array([0.5,0.5]))
```
where `mats` and `vecs` are defined as follows:

```python
mat0 = np.array([[1.0,0.0],[0.0,1.0]])/3
mat1 = np.array([[np.cos(np.pi/3),-np.sin(np.pi/3)],[np.sin(np.pi/3),np.cos(np.pi/3)]])/3
mat2 = np.array([[np.cos(-np.pi/3),-np.sin(-np.pi/3)],[np.sin(-np.pi/3),np.cos(-np.pi/3)]])/3
mat3 = np.array([[1.0,0.0],[0.0,1.0]])/3

mats = np.array([mat0,mat1,mat2,mat3])

vec0 = np.array([0.0,0.0])
vec1 = np.array([1.0/3,0.0])
vec2 = np.array([0.5,np.sqrt(3)/6])
vec3 = np.array([2.0/3,0.0])

vecs = np.array([vec0,vec1,vec2,vec3])
```

<!-- grade_id: cell-908b7a8446becda5, points: 0, grade: false, solution: true -->
```python
def ifs_step(matrices, vectors, coords):
    ### BEGIN SOLUTION
    from numpy.random import randint
    from numpy import dot
    n = len(matrices)
    r = randint(n)
    return dot(matrices[r,:,:],coords)+vectors[r,:]
    ### END SOLUTION
```

<!-- grade_id: cell-cbbb42020161863b, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4a(_globals)
```

Output:
```
Test passed!
1
```

(b) Write, and test on the same matrices and vectors as in part (a), a function called `iterate_ifs` which takes as its arguments `matrices`, `vectors`, an initial value `coords0` and non-negative ints `nskip` and `niterate`. It should then:

<ul>
    <li>iterate <code>ifs_step</code>, using these values of <code>matrices</code> and <code>vectors</code>, a total of <code>nskip</code> times, starting with <code>coords0</code>, and discarding all iterates except the most recent;</li>
    <li>initialize <code>coords_array</code> as a $(1\times2)$ array whose first and only row is the most recent iterate of <code>ifs_step</code>;</li>    
    <li>iterate <code>ifs_step</code>, using these values of <code>matrices</code> and <code>vectors</code>, a total of <code>niterate</code> times, starting with that most recent iterate, and making each fresh iterate a new row of <code>coords_array</code>;</li>
    <li>return the final value of <code>coords_array</code>, which should be a 2D array with dimensions <code>(niterate + 1)</code> by $2$.</li>
</ul>

<!-- grade_id: cell-5d14c0bc80e91162, points: 0, grade: false, solution: true -->
```python
def iterate_ifs(matrices, vectors, coords0, nskip, niterate):
    ### BEGIN SOLUTION
    from numpy import append
    coords = coords0
    for n in range(nskip):
        coords = ifs_step(matrices, vectors, coords)
    coords_array = np.array([coords])
    for n in range(niterate):
        coords = ifs_step(matrices, vectors, coords)
        coords_array = np.append(coords_array,[coords],axis=0)
    return coords_array
    ### END SOLUTION
```

<!-- grade_id: cell-fe82c19cc4286df2, points: 1, grade: true, solution: false -->
```python
# 1 Mark
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4b(_globals)
```

(c) Write a function called `ifs_plot` which should take the same arguments as `iterate_ifs`, together with a keyword-only argument, `markersize`, with default value 0.5. It should then calculate the value of `iterate_ifs(matrices, vectors, coords0, nskip, niterate)`, as a 2D array with dimensions <code>(niterate + 1)</code> by $2$, and create a point plot using column 0 of this array for $x$-values, and column 1 for $y$-values. The marker size should be determined by the value of the keyword-only argument `markersize`. The figure size should be something sensible like 10 by 10, and the scales on the axes should be equal.

Test your function using the matrices and vectors from part (a), with `nskip` set to 100 and `niterate` set to 10000. Choose a suitable marker size. Make the initial coordinates $(0.5, 0.5)$. Make the scales the same on both axes.

Autograding will be no help here, but you ought to see quite an interesting plot.

<!-- grade_id: cell-bc580958691bcd60, points: 0, grade: false, solution: true -->
```python
def ifs_plot(matrices, vectors, coords0, nskip, niterate, *, markersize=0.5):
    ### BEGIN SOLUTION
    import matplotlib.pyplot as plt
    coords_array = iterate_ifs(matrices, vectors, coords0, nskip, niterate)
    plt.plot(coords_array[:,0],coords_array[:,1], '.', markersize=markersize)
ifs_plot(mats,vecs,np.array([0.5,0.5]), 100, 10000, markersize=0.1)
plt.axis('equal');
### END SOLUTION
```

(d) Amend your `ifs_step` function so that it accepts an optional, keyword-only argument, called `weights`, whose default value `None`. If the value of `weights` is `None`, your function should do exactly the same as before (though not necessarily in exactly the same way). 

But optionally, the user may give `weights` a <em>list</em> as its value, such as `[1, 4, 2, 2]`. If she does this, the choice of matrix and vector is made with <em>weighted probabilities</em>; in this case $1/9$, $4/9$, $2/9$ and $2/9$ respectively.

Test your new function.

<!-- grade_id: cell-331cb1647fe3fa18, points: 0, grade: false, solution: true -->
```python
def ifs_step(matrices, vectors, coords, *, weights=None):
    ### BEGIN SOLUTION
    from random import choices
    from numpy import dot
    n = len(matrices)
    r = choices(range(n),weights=weights)[0]
    return dot(matrices[r,:,:],coords)+vectors[r,:]
    ### END SOLUTION
```

<!-- grade_id: cell-1bba6fa67d2c6e91, points: 3, grade: true, solution: false -->
```python
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4d(_globals)
```

(e) Amend your `iterate_ifs` and `ifs_plot` functions to accept and use the optional keyword-only argument `weights`, and test the latter using the transformation matrices

$$\left(\begin{array}{cc}
0.0&0.0\\0.0&0.16\end{array}\right),
\left(\begin{array}{cc}
0.85&0.04\\-0.04&0.85\end{array}\right),
\left(\begin{array}{cc}
0.20&-0.26\\0.23&0.22\end{array}\right),
\left(\begin{array}{cc}
-0.15&0.28\\0.26&0.24\end{array}\right)$$

and the translation vectors

$$\left(\begin{array}{c}
0.0\\0.0\end{array}\right),
\left(\begin{array}{c}
0.0\\1.6\end{array}\right),
\left(\begin{array}{c}
0.0\\1.6\end{array}\right),
\left(\begin{array}{c}
0.0\\0.44\end{array}\right).$$

The weights should be 1, 85, 7 and 7.

Use an `nskip` value of 100, an `niterate` value of 50000 if you can (less if that takes ages). Again, the initial coordinates $(0.5, 0.5)$ should be fine.

Autograding won't help with the final plot, but it ought to look really good.

<!-- grade_id: cell-9da17aa155549a98, points: 0, grade: false, solution: true -->
```python
def iterate_ifs(matrices, vectors, coords0, nskip, niterate, *, weights=None):
    ### BEGIN SOLUTIONx
    from numpy import append
    for n in range(nskip):
        coords0 = ifs_step(matrices, vectors, coords0, weights=weights)
    coords_array = np.array([coords])
    for n in range(niterate):
        coords0 = ifs_step(matrices, vectors, coords0, weights=weights)
        coords_array = np.append(coords_array,[coords0],axis=0)
    return coords_array
    ### END SOLUTION
```

<!-- grade_id: cell-bf0a82655e5faa3e, points: 3, grade: true, solution: false -->
```python
# For testing iterate_ifs
# 3 Marks
# Do not try to delete this cell
# Run this cell for grading
_globals = globals()
question4e(_globals)
```

<!-- grade_id: cell-f3911287c3952ecb, points: 0, grade: false, solution: true -->
```python
def ifs_plot(matrices, vectors, coords0, nskip, niterate, *, markersize=0.5, weights=None):
    ### BEGIN SOLUTION
    import matplotlib.pyplot as plt
    coords_array = iterate_ifs(matrices, vectors, coords0, nskip, niterate, weights=weights)
    plt.plot(coords_array[:,0],coords_array[:,1], '.', markersize=markersize)
mat1 = np.array([[0.0,0.0],[0.0,0.16]])
mat2 = np.array([[0.85,0.04],[-0.04,0.85]])
mat3 = np.array([[0.20,-0.26],[0.23,0.22]])
mat4 = np.array([[-0.15,0.28],[0.26,0.24]])
mats = np.array([mat1,mat2,mat3,mat4])
vec1 = np.array([0.0,0.0])
vec2 = np.array([0.0,1.6])
vec3 = np.array([0.0,1.6])
vec4 = np.array([0.0,0.44])
vecs = np.array([vec1,vec2,vec3,vec4])
ifs_plot(mats,vecs,np.array([0.5,0.5]), 100, 50000, markersize=0.1, 
         weights = [0.01, 0.85, 0.07, 0.07])
plt.axis('equal');
### END SOLUTION
```

## Bonus Question

Consider a list of integers denoting the positions of people queuing in a line, where the person in the front of the queue is integer 1, the second person 2, all the way to the last person in position $n$. In this queue, an individual can bribe the person directly in front of them and can at most make two bribes. Write a function that takes in a list of integers and determines the number of bribes required to get the queue into its current state. If the queue is not possible to achieve through the bribing mechanism described above, then the function should return the string `not possible`. 

**Example:**

Given the input `[2, 1, 5, 3, 4]`, the function should return 3, this being the number of bribes to move from `[1, 2, 3, 4, 5]` to the given input.

<!-- grade_id: cell-f1854e2dfb7cec21, points: 0, grade: false, solution: true -->
```python
def queue_bribe(queue):
    """
    Given an arrangement of the integers between 1 and n,
    returns the number of transpositions needed to obtain
    this arrangement, starting with [1, 2, ..., n], where
    each element may migrate to the left a maximum of two
    places.
    """
    ### BEGIN SOLUTION
    
    # initialize count
    count = 0
    
    # loop over right elements (starting at zeroth!)
    for j in range(len(queue)):
        # pick out element number j
        right_el = queue[j]
        # if right element has migrated more than two places left...
        if right_el - j > 3:
            # ... arrangement is impossible given the rules
            return 'not possible'
        # loop over left elements, going back no further than necessary
        for i in range(max(right_el - 2, 0), j):
            # if there has been a transposition...
            if queue[i] > right_el:
                #... increment the count
                count += 1
    # return total number of transpositions
    return count
    
    ### END SOLUTION
```

<!-- grade_id: cell-c120655e8ffe380d, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
q = [2, 1, 5, 3, 4]
tans = 3
assert(queue_bribe(q) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-11db81758bc1a9a2, points: 1, grade: true, solution: false -->
```python
# Run this test cell to check your code
# Do not delete this cell
# 1 mark
# Unit test
q = [2, 5, 1, 3, 4]
tans = 'not possible'
assert(queue_bribe(q) == tans)
print('Test case passed!!!')
```

<!-- grade_id: cell-64aac9242a875182, points: 3, grade: true, solution: false -->
```python
# Hidden test
# No output will be produced
# 3 marks
### BEGIN HIDDEN TESTS
q = [1, 2, 5, 3, 7, 8, 6, 4]
tans = 7
assert(queue_bribe(q) == tans)
q = [1, 2, 5, 3, 4, 7, 8, 6]
tans = 4
assert(queue_bribe(q) == tans)
print('Hidden test case passed!!!')
### END HIDDEN TESTS
```

```python

```